Linking the evolution of two prefrontal brain regions to social and foraging challenges in primates

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

The present study provides a phylogenetic analysis of the size prefrontal areas in primates, aiming to investigate whether relative size of the rostral prefrontal cortex (frontal pole) and dorsolateral prefrontal cortex volume vary according to known ecological or social variables.

I am very much in favor of the general approach taken in this study. Neuroimaging now allows us to obtain more detailed anatomical data in a much larger range of species than ever before and this study shows the questions that can be asked using these types of data. In general, the study is conducted with care, focusing on anatomical precision in definition of the cortical areas and using appropriate statistical techniques, such as PGLS.

I have read the revised version of the manuscript with interest. I agree with the authors that a focus on ecological vs laboratory variables is a good one, although it might have been useful to reflect that in the title.

I am happy to see that the authors included additional analyses using different definitions of FP and DLPFC in the supplementary material. As I said in my earlier review, the precise delineation of the areas will always be an issue of debate in studies like this, so showing the effects of different decisions in vital.

We thank the reviewer for these positive remarks and for these very useful suggestions on the previous version of this article.

I am sorry the authors are so dismissive of the idea of looking the models where brain size and area size are directly compared in the model, rather preferring to run separate models on brain size and area size. This seems to me a sensible suggestion.

We agree with the reviewer 1 and the response of reviewer 3 also made it clear to us of why it was an important issue. We have therefore addressed it more thoroughly this time.

First, we have added a new analysis, with whole brain volume included as covariate in the model accounting for regional volumes, together with the socio-ecological variables of interest. As expected given the very strong correlation across all brain measures (>90%), the effects of all socio-ecological factors disappear for both FP and DLPFC volumes when ‘whole brain’ is included as covariate. This is coherent with our previous analysis showing that the same combination of socio-ecological variables could account for the volume of FP, DLPFC and the whole brain. Nevertheless, the interpretation of these results remains difficult, because of the hidden assumptions underlying the analysis (see below).

Second, we have clarified the theoretical reasons that made us choose absolute vs relative measures of brain volumes. In short, we understand the notion of specificity associated with relative measures, but 1) the interpretation of relative measures is confusing and 2) we have alternative ways to evaluate the specificity of the effects (which are complementary to the idea of adding whole brain volume as covariate).

Our goal here was to evaluate the influence of socio-ecological factors on specific brain regions, based on their known cognitive functions in laboratory conditions (working memory for the DLPFC and metacognition for the frontal pole). Thus, the null hypothesis is that socio-ecological challenges supposed to mobilize working memory and metacognition do not affect the size of the brain regions associated with these functions (respectively DLPFC and FP). This is what our analysis is testing, and from that perspective, it seems to us that direct measures are better, because within regions (across species), volumes provide a good index of neural counts (since densities are conserved), which are indicative fo the amount of computational resources available for the region. It is not the case when using relative measures, or when using the whole brain as covariate, since densities are heterogenous across brain regions (e.g. Herculano-Houzel, 2011; 2017, but see below for further details on this).

Quantitatively, the theoretical level of specificity of the relation between brain regions and socio-ecological factors is difficult to evaluate, given that our predictions are based on the cognitive functions associated with DLPFC and FP, namely working memory and metacognition, and that each of these cognitive functions also involved other brain regions. We would actually predict that other brain regions associated with the same cognitive functions as DLPFC or FP also show a positive influence of the same socioecological variables. Given that the functional mapping of cognitive functions in the brain remains debated, it is extremely difficult to evaluate quantitatively how specific the influence of the socio-ecological factors should be on DLPFC and FP compared to the rest of the brain, in the frame of our hypothesis.

Critically, given that FP and DLPFC show a differential sensitivity to population density, a proxy for social complexity, and that this difference is in line with laboratory studies showing a stronger implication of the FP in social cognition, we believe that there is indeed some specificity in the relation between specific regions of the PFC and socioecological variables. Thus, our results as a whole seem to indicate that the relation between prefrontal cortex regions and socio-ecological variables shows a small but significant level of specificity. We hope that the addition of the new analysis and the corresponding modifications of the introduction and discussion section will clarify this point.

Similarly, the debate about whether area volume and number of neurons can be equated across the regions is an important one, of which they are a bit dismissive.

We are sorry that the reviewer found us a bit dismissive on this issue, and there may have been a misunderstanding.

Based on the literature, it is clearly established that for a given brain region, area volume provides a good proxy for the number of neurons, and it is legitimate to generalize this relation across species if neuronal densities are conserved for the region of interest (see for example Herculano-Houzel 2011, 2017 for review). It seems to be the case across primates because cytoarchitectonic maps are conserved for FP and DLPFC, at least in humans and laboratory primates (Petrides et al, 2012; Sallet et al, 2013; Gabi et al, 2016; Amiez et al, 2019). But we make no claim about the difference in number of neurons between FP and DLPFC, and we never compared regional volumes across regions (we only compared the influence of socio-ecological factors on each regional volume), so their difference in cellular density is not relevant here. As long as the neuronal density is conserved across species but within a region (DLPFC or FP), the difference in volume for that region, across species, does provide a reliable proxy for the influence of the socioecological regressor of interest (across species) on the number of neurons in that region.

Our claims are based on the strength of the relation between 1) cross-species variability in a set of socio-ecological variables and 2) cross-species variability in neural counts in each region of interest (FP or DLPFC). Since the effects of interest relate to inter-specific differences, within a region, our only assumption is that the neural densities are conserved across distinct species for a given brain region. Again (see previous paragraph), there is reasonable evidence for that in the literature. Given that assumption, regional volumes (across species, for a given brain region) provide a good proxy for the number of neurons. Thus, the influence of a given socio-ecological variable on the interspecific differences in the volume of a single brain region provides a reliable estimate of the influence of that socio-ecological variable on the number of neurons in that region (across species), and potentially of the importance of the cognitive function associated with that region in laboratory conditions. None of our conclusions are based on direct comparison of volumes across regions, and we only compared the influence of socioecological factors (beta weights, after normalization of the variables).

Note that this is yet another reason for not using relative measures and not including whole brain as covariate in the regression model: Given that whole brain and any specific region have a clear difference in density, and that this difference is probably not conserved across species, relative measures (or covariate analysis) cannot be used as proxies for neuronal counts (e.g. Herculano-Houzel, 2011). In other words, using the whole brain to rescale individual brain regions relies upon the assumption that the ratios of volumes (specific region/whole brain) are equivalent to the ratios of neural counts, which is not valid given the differences in densities.

Nevertheless, I think this is an important study. I am happy that we are using imaging data to answer more wider phylogenetic questions. Combining detailed anatomy, big data, and phylogenetic statistical frameworks is a important approach.

We really thank the reviewer for these positive remarks, and we hope that this study will indeed stimulate others using a similar approach.

Reviewer #2 (Public Review):

In the manuscript entitled "Linking the evolution of two prefrontal brain regions to social and foraging challenges in primates" the authors measure the volume of the frontal pole (FP, related to metacognition) and the dorsolateral prefrontal cortex (DLPFC, related to working memory) in 16 primate species to evaluate the influence of socio-ecological factors on the size of these cortical regions. The authors select 11 socio-ecological variables and use a phylogenetic generalized least squares (PGLS) approach to evaluate the joint influence of these socio-ecological variables on the neuro-anatomical variability of FP and DLPFC across the 16 selected primate species; in this way, the authors take into account the phylogenetic relations across primate species in their attempt to discover the the influence of socio-ecological variables on FP and DLPF evolution.

The authors run their studies on brains collected from 1920 to 1970 and preserved in formalin solution. Also, they obtained data from the Mussée National d´Histoire Naturelle in Paris and from the Allen Brain Institute in California. The main findings consist in showing that the volume of the FP, the DLPFC, and the Rest of the Brain (ROB) across the 16 selected primate species is related to three socio-ecological variables: body mass, daily traveled distance, and population density. The authors conclude that metacognition and working memory are critical for foraging in primates and that FP volume is more sensitive to social constraints than DLPFC volume.

The topic addressed in the present manuscript is relevant for understanding human brain evolution from the point of view of primate research, which, unfortunately, is a shrinking field in neuroscience. But the experimental design has two major weak points: the absence of lissencephalic primates among the selected species and the delimitation of FP and DLPFC. Also, a general theoretical and experimental frame linking evolution (phylogeny) and development (ontogeny) is lacking.

We are sorry that the reviewer still believes that these two points are major weaknesses.

- We have added a point on lissencephalic species in the discussion. In short, we acknowledge that our work may not be applied to lissencephalic species because they cannot be studied with our method, but on the other hand, based on laboratory data there is no evidence showing that the functional organization of the DLPFC and FP in lissencephalic primates is radically different from that of other primates (Dias et al, 1996; Roberts et al, 2007; Dureux et al, 2023; Wong et al, 2023). Therefore, there is no a priori reason to believe that not including lissencephalic primates prevents us from drawing conclusions that are valid for primates in general. Moreover, as explained in the discussion, including lissencephalic primates would require using invasive functional studies, only possible in laboratory conditions, which would not be compatible with the number of species (>15) necessary for phylogenetic studies (in particular PGLS approaches). Finally, as pointed out by the reviewer, our study is also relevant for understanding human brain evolution, and as such, including lissencephalic species should not be critical to this understanding.

- In response to the remarks of reviewer 1 on the first version of the manuscript, we had included a new analysis in the previous version of the manuscript, to evaluate the validity of our functional maps given another set of boundaries between FP and DLPFC. But one should keep in mind that our objective here is not to provide a definitive definition of what the regions usually referred to as DLPFC and FP should be from an anatomical point of view. Rather, as our study aims at taking into account the phylogenetic relations across primate species, we chose landmarks that enable a comparison of the volume of cortex involved in metacognition (FP) and working memory (DLPFC) across species. We have also updated the discussion accordingly.

We agree that this is a difficult point and we have always acknowledged that this was a clear limitation in our study. In the light of the functional imaging literature in humans and non-human primates, as well as the neurophysiological data in macaques, defining the functional boundary between FP and DLPFC remains a challenging issue even in very well controlled laboratory conditions. As mentioned by reviewer 1, “the precise delineation of the areas will always be an issue of debate in studies like this, so showing the effects of different decisions in vital”. Again, an additional analyses using different boundaries for FP and DLPFC was included in the supplementary material to address that issue. Now, we are not aware of solid evidence showing that the boundaries that we chose for DLPFC vs FP were wrong, and we believe that the comparison between 2 sets of measures as well as the discussion on this topic should be sufficient for the reader to assess both the strength and the limits of our conclusion. That being said, if the reviewer has any reference in mind showing better ways to delineate the functional boundary between FP and DLPFC in primates, we would be happy to include it in our manuscript.

- The question of development, which is an important question per se, is neither part of the hypothesis nor central for the field of comparative cognition in primates. Indeed, major studies in the field do not mention development (e.g. Byrne, 2000; Kaas, 2012; Barton, 2012). De Casien et al (2022) even showed that developmental constraints are largely irrelevant (see Claim 4 of their article): [« The functional constraints hypothesis […] predicts more complex, ‘mosaic’ patterns of change at the network level, since brain structure should evolve adaptively and in response to changing environments. It also suggests that ‘concerted’ patterns of brain evolution do not represent conclusive evidence for developmental constraints, since allometric relationships between developmentally linked or unlinked brain areas may result from selection to maintain functional connectivity. This is supported by recent computational modeling work [81], which also suggests that the value of mosaic or concerted patterns may fluctuate through time in a variable environment and that developmental coupling may not be a strong evolutionary constraint. Hence, the concept of concerted evolution can be decoupled from that of developmental constraints »].

Finally, when studies on brain evolution and cognition mention development, it is generally to discuss energetic constraints rather than developmental mechanisms per se (Heldstab et al 2022 ; Smaers et al, 2021; Preuss & Wise, 2021; Dunbar & Schutz, 2017; MacLean et al, 2012. Mars et al, 2018; 2021). Therefore, development does not seem to be a critical issue, neither for our article nor for the field.

Reviewer #3 (Public Review):

This is an interesting manuscript that addresses a longstanding debate in evolutionary biology - whether social or ecological factors are primarily responsible for the evolution of the large human brain. To address this, the authors examine the relationship between the size of two prefrontal regions involved in metacognition and working memory (DLPFC and FP) and socioecological variables across 16 primate species. I recommend major revisions to this manuscript due to: 1) a lack of clarity surrounding model construction; and 2) an inappropriate treatment of the relative importance of different predictors (due to a lack of scaling/normalization of predictor variables prior to analysis).

We thank the reviewer for his/her remarks, and for the clarification of his /her criticism regarding the use of relative measures. We are sorry to have missed the importance of this point in the first place. We also thank the reviewer for the cited references, which were very interesting and which we have included in the discussion. As the reviewer 1 also shared these concerns, we wrote a detailed response to explain how we addressed the issue above.

First, we did run a supplementary analysis where whole brain volume was added as covariate, together with socio-ecological variables, to account for the volume of FP or DLPFC. As expected given the very high correlation across all 3 brain measures, none of the socio-ecological variables remained significant. We have added a long paragraph in the discussion to tackle that issue. In short, we agree with the reviewer that the specificity of the effects (on a given brain region vs the rest of the brain) is a critical issue, and we acknowledge that since this is a standard in the field, it was necessary to address the issue and run this extra-analysis. But we also believe that specificity could be assessed by other means: given the differential influence of ‘population density’ on FP and DLPFC, in line with laboratory data, we believe that some of the effects that we describe do show specificity. Also, we prefer absolute measures to relative measures because they provide a better estimate of the corresponding cognitive operation, because standard allometric rules (i.e., body size or whole brain scaling) may not apply to the scaling and evolution of FP and DLPFC in primates.. Indeed, given that we use these measures as proxies of functions (metacognition for FP and working memory for DLPFC), it is clear that other parts of the brain should show the same effect since these functions are supported by entire networks that include not only our regions of interest but also other cortical areas in the parietal lobe. Thus, the extent to which the relation with socio-ecological variables should be stronger in regions of interest vs the whole brain depends upon the extent to which other regions are involved in the same cognitive function as our regions of interest, and this is clearly beyond the scope of this study. More importantly, volumetric measures are taken as proxies for the number of neurons, but this is only valid when comparing data from the same brain region (across species), but not across brain regions, since neural densities are not conserved. Thus, using relative measures (scaling with the whole brain volume) would only work if densities were conserved across brain regions, but it is not the case. From that perspective, the interpretation of absolute measures seems more straightforward, and we hope that the specificity of the effects could be evaluated using the comparison between the 3 measures (FP, DLPFC and whole brain) as well as the analysis suggested by the reviewer. We hope that the additional analysis and the updated discussion will be sufficient to cover that question, and that the reader will have all the information necessary to evaluate the level of specificity and the extent to which our findings can be interpreted.

Recommendations for the authors:

Reviewer #2 (Recommendations For The Authors):

In my previous review of the present manuscript, I pointed out the fact that defining parts, modules, or regions of the primate cerebral cortex based on macroscopic landmarks across primate species is problematic because it prevents comparisons between gyrencephalic and lissencephalic primate species. The authors have rephrased several paragraphs in their manuscript to acknowledge that their findings do apply to gyrencephalic primates.

I also said that "Contemporary developmental biology has showed that the selection of morphological brain features happens within severe developmental constrains. Thus, the authors need a hypothesis linking the evolutionary expansion of FP and DLPFC during development. Otherwise, the claims form the mosaic brain and modularity lack fundamental support". I insisted that the author should clarify their concept of homology of cerebral cortex parts, modules, or regions cross species (in the present manuscript, the frontal pole and the dorsolateral prefrontal cortex). Those are not trivial questions because any phylogenetic explanation of brain region expansion in contemporary phylogenetic and evolutionary biology must be rooted in evolutionary developmental biology. In this regard, the authors could have discussed their findings in the frame of contemporary studies of cerebral cortex evolution and development, but, instead, they have rejected my criticism just saying that they are "not relevant here" or "clearly beyond the scope of this paper".

The question of development, which is an important question per se, is neither part of the hypothesis nor central for the field of comparative cognition in primates. Indeed, the major studies in the field do not mention development and some even showed that developmental constraints were not relevant (see De Casien et al., 2022 and details in our response to the public review). When studies on brain evolution and cognition mention development, it is generally to discuss energetic constraints rather than developmental mechanisms per se (Heldstab et al 2022 ; Smaers et al, 2021; Preuss & Wise, 2021; Dunbar & Schutz, 2017; MacLean et al, 2012. Mars et al, 2018; 2021).

If the other reviewers agree, the authors are free to publish in eLife their correlations in a vacuum of evolutionary developmental biology interpretation. I just disagree. Explanations of neural circuit evolution in primates and other mammalian species should tend to standards like the review in this link: https://royalsocietypublishing.org/doi/full/10.1098/ rstb.2020.0522

In this article, Paul Cizek (a brilliant neurophysiologist) speculates on potential evolutionary mechanisms for some primate brain functions, but there is surprisingly very little reference to the existing literature on primate evolution and cognition. There is virtually no mention of studies that involve a large enough number of species to address evolutionary processes and/or a comparison with fossils and/or an evaluation of specific socio-ecological evolutionary constraints. Most of the cited literature refers to laboratory studies on brain anatomy of a handful of species, and their relevance for evolution remains to be evaluated. These ideas are very interesting and they could definitely provide an original perspective on evolution, but they are mostly based on speculations from laboratory studies, rather than from extensive comparative studies. This paper is interesting for understanding developmental mechanisms and their constraints on neurophysiological processes in laboratory conditions, but we do not think that it would fit it in the framework of our paper as it goes far beyond our main topic.

Reviewer #3 (Recommendations For The Authors):

Yes, I am suggesting that the authors also include analyses with brain size (rather than body size) as a covariate to evaluate the effects of other variables in the model over and above the effect on brain size. In a very simplified theoretical scenario: two species have the same body sizes, but species A has a larger brain and therefore a larger FP. In this case, species A has a larger FP because of brain allometric patterns, and models including body size as a covariate would link FP size and socioecological variables characteristic of species A (and others like it). However, perhaps the FP of species A is actually smaller than expected for its brain size, while the FP of species B is larger than expected for its brain size.

As explained in our response to the public review, we did run this analysis and we agree with the reviewer’s point from a practical point of view: it is important to know the extent to which the relation with a set of socio-ecological variables is specific of the region of interest, vs less specific and present for other brain regions. Again, we are sorry to not have understood that earlier, and we acknowledge that since it is a standard in the field, it needs to be addressed thoroughly.

We understand that the scaling intuition, and the need to get a reference point for volumetric measures, but here the volume of each brain region is taken as a proxy for the number of neurons and therefore for the region’s computational capacities. Since, for a given brain region (FP or DLPFC) the neural densities seem to be well conserved across species, comparing regional volumes across species provides a good proxy for the contrast (across species) in neural counts for that region. All we predicted was that for a given brain region, associated with a given cognitive operation, the volume (number of neurons) would be greater in species for which socio-ecological constraints potentially involving that specific cognitive operation were greater. We do not understand how or why the rest of the brain would change this interpretation (of course, as discussed just above, beyond the question of specificity). And using whole brain volume as a scaling measure is problematic because the whole brain density is very different from the density of these regions of the prefrontal cortex (see above for further details). Again, we acknowledge that allometric patterns exist, and we understand how they can be interpreted, but we do not understand how it could prove or disprove our hypothesis (brain regions involved in specific cognitive operations are influenced by a specific set of socio-ecological variables). When using volumes as a proxy for computational capacities, the theoretical implications of scaling procedures might be problematic. For example, it implies that the computational capacities of a given brain region are scaled by the rest of the brain. All other things being equal, the computational capacities of a given brain region, taken as the number of neurons, should decrease when the size of the rest of the brain increases. But to our knowledge there is no evidence for that in the literature. Clearly these are very challenging issues, and our position was to take absolute measures because they do not rely upon hidden assumptions regarding allometric relations and their consequence on cognition.

But since we definitely understand that scaling is a reference in the field, we have not only completed the corresponding analysis (including the whole brain as a covariate, together with socio-ecological variables) but also expended the discussion to address this issue in detail. We hope that between this new analysis and the comparison of effects between non-scaled measures of FP, DLPFC and the whole brain, the reader will be able to judge the specificity of the effect.

Models including brain (instead of body) size would instead link FP size and socioecological variables characteristic of species B (and others like it). This approach is supported by a large body of literature linking comparative variation in the relative size of specific brain regions (i.e., relative to brain size) to behavioral variation across species - e.g., relative size of visual/olfactory brain areas and diurnality/nocturnality in primates (Barton et al. 1995), relative size of the hippocampus and food caching in birds (Krebs et al. 1989).

Barton, R., Purvis, A., & Harvey, P. H. (1995). Evolutionary radiation of visual and olfactory brain systems in primates, bats and insectivores. Philosophical Transactions of the Royal Society of London. Series B: Biological Sciences, 348(1326), 381-392.

Krebs, J. R., Sherry, D. F., Healy, S. D., Perry, V. H., & Vaccarino, A. L. (1989). Hippocampal specialization of food-storing birds. Proceedings of the National Academy of Sciences, 86(4), 1388-1392.

We are grateful to the reviewer for mentioning these very interesting articles, and more generally for helping us to understand this issue and clarify the related discussion. Again, we understand the scaling principle but the fact that these methods provide interesting results does not make other approaches (such as ours) wrong or irrelevant. Since we have used both our original approach and the standard version as requested by the reviewer, the reader should be able to get a clear picture of the measures and of their theoretical implications. We sincerely hope that the present version of the paper will be satisfactory, not only because it is clearer, but also because it might stimulate further discussion on this complex question.

Changes from the earlier version to the revised version

Abstract

Abstract

The diversity of cognitive skills across primates remains both a fascinating and a controversial issue. Recent comparative studies provided conflicting results regarding the contribution of social vs ecological constraints to the evolution of cognition. Here, we used an interdisciplinary approach combining comparative cognitive neurosciences and behavioral ecology. Using brain imaging data from 16 primate species, we measured the size of two prefrontal brain regions, the frontal pole (FP) and the dorso-lateral prefrontal cortex (DLPFC), respectively involved in metacognition and working memory, and examined their relation to a combination of socio-ecological variables. The size of these prefrontal regions, as well as the whole brain, was best explained by three variables: body mass, daily traveled distance (an index of ecological constraints) and population density (an index of social constraints). The strong influence of ecological constraints on FP and DLPFC volumes suggests that both metacognition and working memory are critical for foraging in primates. Interestingly, FP volume was much more sensitive to social constraints than DLPFC volume, in line with laboratory studies showing an implication of FP in complex social interactions. Thus, our data highlights the relative weight of social vs ecological constraints on the evolution of specific prefrontal brain regions and their associated cognitive operations in primates.

Introduction

Earlier version: Numerous studies have addressed the mechanisms underlying the evolution of cognitive abilities in primates, using brain size as a proxy (Powell, Isler et al. 2017, Chambers, Heldstab et al. 2021, van Schaik, Triki et al. 2021, DeCasien, Barton et al. 2022). One of the major hypotheses, referred to as the ‘social brain’ hypothesis, proposes that social interactions require higher cognitive skills such that the complexity of social interactions played a central role in the increase in brain size during primate evolution (Dunbar 1998). Another major theory, referred to as the ‘ecological brain’ hypothesis, proposes that the increase in brain size and corresponding cognitive skills was driven by the need to build mental representations in order to forage efficiently (Milton 1981, Milton 1993). Critically, both of these theories are based on the assumption that the whole brain volume is a good proxy for overall cognitive abilities, and that these global abilities can benefit social interactions or foraging, across species in an evolutionary framework. But clearly, the level of precision of such neuro-cognitive scenario remains limited, given the precision with which neuroscience and psychology have characterized the brain functional anatomy and cognition, respectively.

Revised version: Numerous studies have addressed the mechanisms underlying the evolution of cognitive abilities in primates, using brain size as a proxy (Chambers, Heldstab, & O’Hara, 2021; A. R. DeCasien, Barton, & Higham, 2022; Powell, Isler, & Barton, 2017; Carel P. van Schaik, Triki, Bshary, & Held-stab, 2021). One of the major hypotheses, referred to as the ‘social brain’ hypothesis, proposes that social interactions require higher cognitive skills such that the complexity of social interactions played a central role in the increase in brain size during primate evolution (Dunbar, 1998). Another major theory, referred to as the ‘ecological brain’ hypothesis, proposes that the increase in brain size and corresponding cognitive skills was driven by the need to build mental representations in order to forage efficiently (Milton, 1981, 1993). Critically, both of these theories are based on the assumption that the whole brain volume is a good proxy for overall cognitive abilities, and that these global abilities can benefit social interactions or foraging, across species in an evolutionary framework. But clearly, the level of precision of such neuro-cognitive scenario remains limited, given the precision with which neuroscience and psychology have characterized the brain functional anatomy and cognition, respectively.

Removed in the revised version: Imaging studies in humans have identified a network of brain regions involved in social interactions, and this network is often referred to as the social brain (Stanley and Adolphs 2013). Recent comparative studies have shown that it also existed in macaques (Sallet, Mars et al. 2011, Rushworth, Mars et al. 2013, Testard, Brent et al. 2022). Along the same lines, laboratory studies have also proposed that a set of specific brain regions (e.g. cingulate cortex) was underlying foraging in humans and macaques (eg (Hayden, Pearson et al. 2011, Kolling, Behrens et al. 2012). These studies, however, have only been conducted in a handful of species and in very artificial conditions, such that the relation between these brain regions and natural behavior or evolution remains elusive. In terms of the cognitive processes involved, ‘social interactions’ and ‘foraging’ remain poorly specific and at least in primates, they probably involve myriads of more elementary cognitive operations. For example, social interactions involve some form of categorization of facial expressions and metacognition, and the brain regions associated with these elementary cognitive operations (e.g. frontal pole or temporo-parietal junction) are usually part of the ‘social brain’, i.e. the set of brain regions associated with social interactions (Gallagher and Frith 2003, Rushworth, Mars et al. 2013, Devaine, San-Galli et al. 2017, de Gelder and Poyo Solanas 2021, de Gelder 2023, Deen, Schwiedrzik et al. 2023). Along the same lines, foraging is thought to require mental spatio-temporal representations of food availability and values, and these elementary cognitive operations rely upon a distinct set of brain regions, including the hippocampus and the ventromedial prefrontal cortex (Zuberbühler and Janmaat 2010, Lin, Horner et al. 2015, Louail, Gilissen et al. 2019, Vikbladh, Meager et al. 2019). Finally, the brain systems potentially involved in social cognition and foraging partially overlap, especially in the frontal lobes (Gallagher and Frith 2003, Rudebeck, Walton et al. 2006, Kolling, Behrens et al. 2012, Yoshida, Saito et al. 2012, Barbey, Colom et al. 2014, Mansouri, Koechlin et al. 2017). This partial overlap between the networks referred to as the ‘social brain’ and the ‘ecological brain’ is in line with the idea that some elementary cognitive operations are generic enough to be involved in multiple contexts (Garcia, Bouret et al. 2021, Shultz and Dunbar 2022). This is typically the case for executive functions (e.g. working memory), which allow to organize behavior over space and time to reach a goal, and rely upon several regions of the prefrontal cortex (Luria 1973, Fuster 2008).

Earlier version: Given the complexity encountered in natural environments, it is hazardous to extrapolate these findings in artificial laboratory conditions and in a handful of species to the decision-making processes in both social and food dimensions occurring in wild primates, for which the diversity of social interactions (e.g. number of individuals an animal can remember, degree of social awareness i.e. knowledge and representation of the dominance hierarchy, kinship relations or friendship associations) and foraging behaviors (e.g. food processing techniques, knowledge of harvesting schedules, extractive foraging) require much more cognitive flexibility. One can wonder to what extent these neuro-cognitive concepts (i.e. elementary cognitive operation and their underlying cerebral substrate) can be used as theoretical bricks to understand the diversity of natural behaviors reported in the wild and their evolution across species.

Revised version: Given the complexity encountered in natural environments, it is hazardous to extrapolate these findings in artificial laboratory conditions to the decision-making processes (in both social and ecological dimensions) occurring in wild primates, for which the diversity of social interactions (e.g. number of individuals an animal can remember, degree of social awareness i.e. knowledge and representation of the dominance hierarchy, kinship relations or friendship associations) and foraging behaviors (e.g. food processing techniques, knowledge of harvesting schedules, extractive foraging) require much more cognitive flexibility. One can wonder to what extent these neurocognitive concepts (i.e. elementary cognitive operation and their underlying cerebral substrate) can be used as theoretical building blocks to understand the diversity of natural behaviors reported in the wild and their evolution across species. In terms of evolution, the functional heterogeneity of distinct brain regions is captured by the notion of ‘mosaic brain’, where distinct brain regions could show a specific relation with various socio-ecological challenges, and therefore have relatively separate evolutionary trajectories (Barton & Harvey, 2000; Alex R. DeCasien & Higham, 2019a). But these relations between specific brain regions and socio-ecological variables across species provide little insight into the cognitive processes that could be at play, because they do not necessarily map onto specific neuro-cognitive processes identified in rigorous laboratory conditions. Thus, it remains challenging to bridge the gap between 1) laboratory studies relating brain regions to elementary cognitive operations in individual species and 2) studies relating brain regions to socio-ecological challenges through an evolutionary scenario (i.e. across numerous species).

Added in the revised version: Imaging studies in humans have identified a network of brain regions involved in social interactions, and this network is often referred to as the social brain (Stanley & Adolphs, 2013). Recent comparative studies have shown that it also existed in macaques (Rushworth, Mars, & Sallet, 2013; Jérôme Sallet et al., 2011; Testard et al., 2022). Along the same lines, laboratory studies have proposed that specific brain regions (e.g. cingulate cortex) were underlying foraging in humans and macaques (e.g. (Hayden, Pearson, & Platt, 2011; Kolling, Behrens, Mars, & Rushworth, 2012). These studies, however, have only been conducted in a handful of species and in very artificial conditions, such that the relation between these brain regions and natural behavior or evolution remains elusive. In terms of the cognitive processes involved, “social interactions” and “foraging” remain poorly specific and they probably involve myriads of more elementary cognitive operations. For example, social interactions involve some form of categorization of facial expressions and metacognition, and the brain regions associated with these elementary cognitive operations (e.g. frontal pole or temporo-parietal junction) are usually part of the “social brain”, i.e. the set of brain regions associated with social interactions (B. de Gelder, 2023; Beatrice de Gelder & Poyo Solanas, 2021; Deen, Schwiedrzik, Sliwa, & Freiwald, 2023; Devaine et al., 2017; Gallagher & Frith, 2003; Rushworth et al., 2013). Along the same lines, foraging is thought to require spatio-temporal mental representations of food availability and values, and these elementary cognitive operations rely upon a distinct set of brain regions, including the hippocampus and the ventrome-dial prefrontal cortex (Lin, Horner, Bisby, & Burgess, 2015; Louail, Gilissen, Prat, Garcia, & Bouret, 2019; Rosati, 2017; Vikbladh et al., 2019; Zuberbühler & Janmaat, 2010). Finally, the brain systems potentially involved in social cognition and foraging partially overlap, especially in the frontal lobes (Barbey et al., 2014; Gallagher & Frith, 2003; Kolling et al., 2012; Mansouri, Koechlin, Rosa, & Buckley, 2017; Rudebeck, Walton, Smyth, Bannerman, & Rushworth, 2006; Yoshida, Saito, Iriki, & Isoda, 2012). This partial overlap between the networks involved in social and foraging cognition laboratory conditions in is in line with the idea that some elementary cognitive operations are generic enough to be involved in multiple contexts (Garcia, Bouret, Druelle, & Prat, 2021; Shultz & Dunbar, 2022). This is typically the case for executive functions (e.g. working memory), which allow to organize behavior over space and time to reach a goal, and rely upon several regions of the prefrontal cortex (Fuster, 2008; Luria, 1973).

Removed in the revised version: In terms of evolution, the functional heterogeneity of distinct brain regions is captured by the notion of ‘mosaic brain’, where distinct brain regions could show a specific relation with various socio-ecological challenges, and therefore have relatively separate evolutionary trajectories (Barton and Harvey 2000, DeCasien and Higham 2019). But these relations between specific brain regions and socio-ecological variables across species provide little insight into the cognitive processes that could be at play. Thus, it remains challenging to bridge the gap between 1) laboratory studies relating brain regions to elementary cognitive operations in individual species and 2) studies relating brain regions to socio-ecological challenges through an evolutionary scenario (i.e. across numerous species).

Earlier version: In order to bridge this gap, we designed this study to derive and test predictions where a set of well-identified elementary cognitive operations and their associated brain regions (as studied in laboratory conditions) could be involved in more natural functions (as assessed in wild animals living in their natural environments). We focused on two well-known cognitive operations, both considered as executive functions, and involving distinct regions of the prefrontal cortex: metacognition, which involves the frontal pole (FP), and working memory, which involves the dorso-lateral prefrontal cortex (DLPFC), as demonstrated by extensive work in humans and macaques (Fuster 2008, Passingham, Passingham et al. 2012, Mansouri, Koechlin et al. 2017). Based on the literature, we hypothesized that FP and metacognition would be associated both with social interactions (by supporting theory of mind, the ability of an individual to conceptualize others’ states of mind) and with foraging (by enabling complex planning) (Frith 2007, Fleming and Dolan 2012, Devaine, San-Galli et al. 2017, Mansouri, Koechlin et al. 2017). For example, group hunting, which involves both social and foraging functions, only occurs in a few primate species where metacognition is thought to be particularly developed (Boesch 1994, van Schaik 2016, Conard, Serangeli et al. 2020, Garcia, Bouret et al. 2021). Concerning the DLPFC, given its very clear implication in working memory and planning, we speculated that it would be critically involved in foraging. Its potential relation with social interactions, however, was less clear: on one hand, DLPFC is rarely associated with social interactions in laboratory studies, but on the other hand working memory and planning could readily be involved in complex social interactions in more natural conditions (Frith 2007, Fuster 2008, Passingham, Passingham et al. 2012, Sallet, Mars et al. 2013).

Revised version: In order to bridge this gap, we designed this study to derive and test predictions where a set of well-identified elementary cognitive operations and their associated brain regions (as studied in laboratory conditions) could be involved in more natural functions (as assessed in wild animals living in their natural environments). We focused on two well-known cognitive operations, both considered as executive functions, and involving distinct regions of the prefrontal cortex: metacognition, which involves the frontal pole (FP), and working memory, which involves the dorso-lateral prefrontal cortex (DLPFC), as demonstrated by extensive work in humans and macaques (Fuster, 2008; Mansouri et al., 2017; Passingham, Passingham, & Wise, 2012). Based on the literature, we hypothesized that FP and metacognition would be associated both with social interactions (by supporting theory of mind, i.e. the ability of an individual to conceptualize others’ states of mind) and with foraging (by enabling complex planning) (Devaine et al., 2017; Fleming & Dolan, 2012; Frith, 2007; Mansouri et al., 2017). For example, group hunting, which involves both social and foraging functions, only occurs in a few primate species where metacognition is thought to be particularly developed (Boesch, 1994; Conard, Serangeli, Bigga, & Rots, 2020; Garcia et al., 2021; C. P. van Schaik, 2016). Concerning the DLPFC, given its very clear implication in working memory and planning, we speculated that it would be critically involved in foraging. Its potential relation with social interactions, however, was less clear: on one hand, DLPFC is rarely associated with social interactions in laboratory studies (Frith, 2007; Fuster, 2008; Passingham et al., 2012; J. Sallet et al., 2013). But, on the other hand working memory and planning could readily be involved in complex social interactions in more natural conditions.

Earlier version: To evaluate the relevance of laboratory-based neuro-cognitive concepts for understanding which cognitive operations could be mobilized to overcome specific socio-ecological challenges in the wild and their variability across species, we addressed two central questions: 1) to what extent does this functional mapping of prefrontal regions generalize across primates, and 2) to what extent does the functional mapping between brain regions and specific executive functions observed in laboratory relate to socio-ecological variables measured in natural conditions?

Revised version: To evaluate the relevance of laboratory-based neuro-cognitive concepts for understanding which cognitive operations could be mobilized to overcome specific socio-ecological challenges in the wild and their variability across species, we addressed the following central question: to what extent does the functional mapping between brain regions and specific executive functions in laboratory conditions relate to variability in socio-ecological variables at the interspecific level? For this purpose, it is necessary to go beyond the handful of species usually studied in well-controlled laboratory conditions, not only to evaluate the reliability of the relations between neuro-anatomical measures and socio-ecological factors across primates, but also to quantify the contribution of phylogeny.

Earlier version: To address these questions, we measured the volume of the two prefrontal regions of interest (FP and DLPFC) in 16 primate species and evaluated the influence of key socio-ecological factors on the size of these regions. Based on the known positive relation between the size of a given brain region (or the number of corresponding neurons), and the relative importance of its associated function, both within and across species, we assume that the size of each brain region can be taken as a proxy for the weight of its associated function on behavior (Barton and Harvey 2000, Maguire, Gadian et al. 2000, Sallet, Mars et al. 2011, Barks, Calhoun et al. 2014, Herculano-Houzel, Kaas et al. 2016, Herculano-Houzel 2017, DeCasien and Higham 2019, Louail, Gilissen et al. 2019, Ferrucci, Nougaret et al. 2022, Testard, Brent et al. 2022). These two regions were chosen and identified based on functional maps and reliable macroscopic landmarks, rather than cytoarchitectonic criteria, to maximize the reliability of the comparative approach (see also (Louail, Gilissen et al. 2019). Anatomical studies have shown that these regions and landmarks were relatively well conserved between humans and macaques, both in terms of cytoarchitectonics and connectivity profiles (Petrides, Tomaiuolo et al. 2012, Sallet, Mars et al. 2013, Amiez, Sallet et al. 2019). Finally, the cognitive functions in which they are involved appear relatively well conserved between macaques and humans (Fuster 2008, Mansouri, Koechlin et al. 2017). In addition to these two regions, we also used measures of whole brain volume as a reference to facilitate comparison with other studies that only use whole brain measures. The 16 primate species were chosen to cover a wide range of phylogenetic distances coupled to a diversity of socio-ecological variables. We compared the influence of several combinations of these socio-ecological variables on the variability in size of each brain region of interest, and interpreted it in the light of its known function in laboratory conditions. Thus, we could evaluate the strength of the relation between specific regions of the prefrontal cortex and specific socio-ecological variables across species, as observed in the wild, thereby complementing laboratory studies and bridging the gap between cognitive neurosciences, behavioral ecology and primate evolution.

Revised version: To address this question, we measured the volume of the two prefrontal regions of interest (FP and DLPFC) in 16 primate species and evaluated the influence of key socio-ecological factors on the size of these regions. Based on the known positive relation between the size of a given brain region (or the number of corresponding neurons), and the relative importance of its associated function, both within and across species, we assume that the size of each brain region can be taken as a proxy for the weight of its associated function on behavior (Barks et al., 2014; Barton & Harvey, 2000; Alex R. DeCasien & Higham, 2019a; Ferrucci et al., 2022; Herculano-Houzel, 2017; Herculano-Houzel, Kaas, & de Oliveira-Souza, 2016; Louail et al., 2019; Maguire et al., 2000; Jé-rôme Sallet et al., 2011; Testard et al., 2022). Anatomical studies have shown that these regions were relatively well conserved between humans and macaques, both in terms of cytoarchitectonics and connectivity profiles (Amiez et al., 2019; Michael Petrides, Tomaiuolo, Yeterian, & Pandya, 2012; J. Sallet et al., 2013). Here, these two regions were chosen and identified based on functional maps and reliable macroscopic landmarks, rather than cytoarchitectonic criteria, to maximize the reliability of the comparative approach (see also (Louail et al., 2019). Finally, the cognitive functions in which they are involved appear relatively well conserved between macaques and humans (Ferrucci et al., 2022; Fuster, 2008; Mansouri et al., 2017; Passingham & Sakai, 2004). The 16 primate species were chosen to cover a wide range of phylogenetic distances coupled to a diversity of socio-ecological variables. We compared the influence of several combinations of these socio-ecological variables on the variability in size of each brain region of interest, and interpreted it in the light of its known function in laboratory conditions. Thus, for a given functional region, the differences in regional volume across species provide a reliable index of the differences in neuronal count, and therefore of the skill level in the corresponding function. Even though comparative studies often used scaled measures (i.e. relative to the whole brain) to account for potential allometric effects, we favored absolute measures because scaling procedures would distort the relation between volume and neural counts (Barton, Purvis, & Harvey, 1995; A. R. DeCasien et al., 2022; Gabi et al., 2016; Herculano-Houzel, 2017; Herculano-Houzel, Collins, Wong, Kaas, & Lent, 2008; Krebs, Sherry, Healy, Perry, & Vaccarino, 1989; Smaers, Gomez-Robles, Parks, & Sherwood, 2017; Smaers & Soligo, 2013). Moreover, the extent to which other brain regions could be affected by socio-ecological variables, and therefore affect scaled measures, is difficult to evaluate. We therefore used several procedures to evaluate the specificity of the effects of socio-ecological factors on FP and DLPFC. First, we compared the influence of socio-ecological variables not only on FP and DLPFC but also on the whole brain. Second, we included the whole brain volume as a covariate along with the socio-ecological variables to account for FP and DLPFC volumes. Altogether, these procedures provide a reliable indication of the specificity of the influence of socio-ecological variables on the FP and DLPFC volumes, respectively used as proxies for inter-species differences in metacognitive and working memory skills. Thus, we could evaluate the strength of the relation between specific regions of the prefrontal cortex and specific socio-ecological variables across species, as observed in the wild, thereby complementing laboratory studies and bridging the gap between cognitive neurosciences, behavioral ecology and primate evolution.

Material & Methods

Sample

Earlier version: Thirty one brain magnetic resonance (MRI) 3D reconstructions from 16 primate species (Ateles fusciceps, n=1; Cebus capucinus, n=1; Cercopithecus mitis, n= 1; Gorilla gorilla, n=5; Gorilla beringei, n=1; Homo sapiens, n=4; Hylobates lar, n=1; Lagothrix lagotricha, n=1; Lophocebus albigena, n=1; Macaca fascicularis, n=1; Macaca fuscata, n=4; Macaca mulatta, n=2; Pan troglodytes, n=5; Pan paniscus, n=1; Papio papio, n=1, Pongo pygmaeus, n=1) were used in this study. Ja-panese macaques (M. fuscata) and rhesus macaques (M. mulatta) were captive animals scanned at the National Institutes for Quantum and Radiological Science and Technology (Chiba, Japan) and at Brain and Spine Institute (Paris, France), respectively. P. troglodytes and G. gorilla brains came from the Muséum national d’Histoire naturelle (Paris, France). They had been collected between 1920 and 1970 and subsequently preserved in formalin solution. The P. paniscus and G. beringei brains came from the Royal Museum for Central Africa (RMCA) (Tervuren, Belgium) and the Royal Belgian Institute of Natural Sciences (Bruxelles, Belgium). All the Pan and Gorilla specimens have been scanned at University of Leuven (KUL). The brain scans A. fusciceps, C. mitis, and L. lagotricha were obtained from the Primate Brain Bank, NIN Utrecht University. Finally, the remaining species (C. capucinus, H. lar, L. albigena, P. pygmaeus, M. fascicularis, P. papio and one specimen of H. sapiens) came from the brain catalogue website (https://braincatalogue.org), which gathers scans of specimens from the collections of the Muséum national d’Histoire naturelle (Paris). The three other H. sapiens brain scans were obtained from the Allen Institute (online brain atlas). All specimens were sexually mature at the time of scanning. The sexes of individuals were mostly unknown. Moreover, some specimens in the sample came from captivity. Thus, we neglected the effects of captivity and sex on brain/endocranium measurements, which were both shown to be very small compared to inter-species differences (Isler and van Schaik 2012).

Revised version: Thirty one brain magnetic resonance (MRI) 3D reconstructions from 16 primate species (Ateles fusciceps, n=1; Cebus capucinus, n=1; Cercopithecus mitis, n= 1; Gorilla gorilla, n=5; Gorilla beringei, n=1; Homo sapiens, n=4; Hylobates lar, n=1; Lagothrix lagotricha, n=1; Lophocebus albigena, n=1; Macaca fascicularis, n=1; Macaca fuscata, n=4; Macaca mulatta, n=2; Pan troglodytes, n=5; Pan paniscus, n=1; Papio papio, n=1, Pongo pygmaeus, n=1) were used in this study. Japanese macaques (M. fuscata) and rhesus macaques (M. mulatta) were captive animals scanned at the National Institutes for Quantum and Radiological Science and Technology (Chiba, Japan) and at Brain and Spine Institute (Paris, France), respectively. P. troglodytes and G. gorilla brains came from the Muséum national d’Histoire naturelle (Paris, France). They had been collected between 1920 and 1970 and subsequently preserved in formalin solution. The P. paniscus and G. beringei brains came from the Royal Museum for Central Africa (RMCA) (Tervuren, Belgium) and the Royal Belgian Institute of Natural Sciences (Bruxelles, Belgium). All the Pan and Gorilla specimens have been scanned at University of Leuven (KUL). The brain scans A. fusciceps, C. mitis, and L. lagotricha were obtained from the Primate Brain Bank, NIN Utrecht University. Finally, the remaining species (C. capucinus, H. lar, L. albigena, P. pygmaeus, M. fascicularis, P. papio and one specimen of H. sapiens) came from the brain catalogue website (https://braincatalogue.org), which gathers scans of specimens from the collections of the Muséum national d’Histoire naturelle (Paris). The three other H. sapiens brain scans were obtained from the Allen Institute (online brain atlas). All specimens were sexually mature at the time of scanning. The sexes of individuals were mostly unknown. Moreover, some specimens in the sample came from captivity. Thus, we neglected the effects of captivity and sex on brain/endocranium measurements, which were both shown to be very small compared to inter-species differences (Isler & van Schaik, 2012).

Processing of brain MRI and measurements

Brain measurements (visualization, segmentation and quantification of brain tissues volumes) were processed using Avizo v9.0 software. The whole brain volume was measured in order to facilitate comparisons with the literature. The cerebellum was excluded from all brain measurements, because it was missing on some of the MRI scans (gorilla brains). Whole brain segmentation was performed using the semi-automated tool in Avizo that enables to select a material or structure according to a specific gray-level threshold. It was however necessary to bring some manual corrections, for example when the brain had a similar gray-level than an adjacent tissue. Segmentations of the frontal pole (FP) and the dorso-lateral prefrontal cortex (DLPFC) were carried out manually with the brush tool, slice by slice of the MRI scan.

Earlier version: The frontal pole is the most rostral part of the prefrontal cortex. Cytoarchitectonic studies indicate that it strongly overlaps with Brodmann area 10 (BA 10) (Semendeferi, Armstrong et al. 2001, Ongür, Ferry et al. 2003, Ramnani and Owen 2004, Tsujimoto, Genovesio et al. 2011). Besides cytoarchitectonic landmarks, the frontal pole can also be identified based on connectivity patterns: it receives projections from the temporal superior cortex (Petrides 2005) and has connections with the superior temporal sulcus (Sallet, Mars et al. 2013). We also used the probabilistic maps of the frontal pole proposed by John et al. (John, Yashavantha et al. 2007) and Bludau et al. (Bludau, Eickhoff et al. 2014). We delimited the frontal pole according to different criteria: it should match the functional anatomy for known species (macaques and humans, essentially) and be reliable enough to be applied to other species using macroscopic neuroanatomical landmarks. We integrated these criteria and the data from the literature on brain atlases of rhesus macaques and humans (Saleem and Logothetis 2007, Borden, Stefan et al. 2015) to define visible limits of the frontal pole, as shown in Figure 1. The anterior limit of the frontal pole was defined as the anterior limit of the brain. The cingulate sulcus represented the posterior limit. The ventral limit was set as the dorsal limit of the gyrus rectus. Finally, the dorsal limit was defined as the fundus of the superior frontal sulcus in humans and apes, or the principal sulcus in monkeys. We chose this very conservative dorsal limit in monkeys to avoid any overlap with what could be considered as part of the DLPFC (typically, BA9). However, we are aware that several studies in macaques suggest that the frontal pole in monkeys might extend more dorsally (Tsujimoto, Genovesio et al. 2011, Sallet, Mars et al. 2013). Thus, we also conducted a series of measures with a less conservative dorsal limit for the frontal pole (see supplementary material section 4).

Revised version: The frontal pole is the most rostral part of the prefrontal cortex. Cytoarchitectonic studies indicate that it strongly overlaps with Brodmann area 10 (BA 10) (Ongür, Ferry, & Price, 2003; Ramnani & Owen, 2004; Semendeferi, Armstrong, Schleicher, Zilles, & Van Hoesen, 2001; Tsujimoto, Genovesio, & Wise, 2011). Besides cytoarchitectonic landmarks, the frontal pole can also be identified based on connectivity patterns: it receives projections from the temporal superior cortex (M. Petrides, 2005) and has connections with the superior temporal sulcus (J. Sallet et al., 2013). We also used the probabilistic maps of the frontal pole proposed by John et al. (John et al., 2007) and Bludau et al. (Bludau et al., 2014). We delimited the frontal pole according to different criteria: it should match the functional anatomy for known species (macaques and humans, essentially) and be reliable enough to be applied to other species using macroscopic neuroanatomical landmarks. We integrated these criteria and the data from the literature on brain atlases of rhesus macaques and humans (Borden, Stefan, & Forseen, 2015; Saleem & Logothetis, 2007) to define visible limits of the frontal pole, as shown in Figure 1. The anterior limit of the frontal pole was defined as the anterior limit of the brain. The cingulate sulcus represented the posterior limit. The ventral limit was set as the dorsal limit of the gyrus rectus. Finally, the dorsal limit was defined as the fundus of the superior frontal sulcus in humans and apes, or the principal sulcus in monkeys. We chose this very conservative dorsal limit in monkeys to avoid any overlap with what could be considered as part of the DLPFC (typically, BA9). However, we are aware that several studies in macaques suggest that the frontal pole in monkeys might extend more dorsally (J. Sallet et al., 2013; Tsujimoto et al., 2011). Thus, we also conducted a series of measures with a less conservative dorsal limit for the frontal pole (see supplementary material section 4).

Boundaries of the frontal pole. From left to right: sagittal view, coronal view for great apes and humans, coronal view for monkeys. Frontal pole is in green. Abbreviations: cs: cingulate sulcus; cc: corpus callosum; rs: rostral sulcus; mos: medial-orbital sulcus; frss: superior frontal sulcus; ps: principal sulcus; D: dorsal; V: ventral; M: medial; L: lateral; A: anterior; P: posterior.

Earlier version: Similarly, the dorso-lateral prefrontal cortex was measured by combining functional and anatomical data from the literature to identify reliable macroscopic landmarks (Levy and Goldman-Rakic 2000, Passingham and Sakai 2004, Passingham, Passingham et al. 2012, Petrides, Tomaiuolo et al. 2012, Sallet, Mars et al. 2013). Given the major difference in sulcal organization between monkeys and greats apes, we used different landmarks, shown on figures 2a and 2b. DLPFC comprises portions of middle frontal gyrus and superior frontal gyrus in great apes and lies in and around the principal sulcus in macaques. Then, the ventral limit was set as the fundus of the principal sulcus for monkeys, and the frontal inferior sulcus for apes and humans. The anterior limit of the DLPFC was defined as the posterior limit of the frontal pole, which was the cingulate sulcus. For apes and humans, the posterior limit was defined as the precentral sulcus, whereas in monkeys it was defined as the end of the arcuate sulcus. Finally, the medial limit was designated as the inter-hemispheric sulcus.

Revised version: Similarly, the dorso-lateral prefrontal cortex was measured by combining functional and anatomical data from the literature to identify reliable macroscopic landmarks (Levy & Goldman-Rakic, 2000; Passingham et al., 2012; Passingham & Sakai, 2004; Michael Petrides et al., 2012; J. Sallet et al., 2013). Given the major difference in sulcal organization between monkeys and greats apes, we used different landmarks, shown on figures 2a and 2b. DLPFC comprises portions of middle frontal gyrus and superior frontal gyrus in great apes and lies in and around the principal sulcus in macaques. Then, the ventral limit was set as the fundus of the principal sulcus for monkeys, and the frontal inferior sulcus for apes and humans. The anterior limit of the DLPFC was defined as the posterior limit of the frontal pole, which was the cingulate sulcus. For apes and humans, the posterior limit was defined as the precentral sulcus, whereas in monkeys it was defined as the end of the arcuate sulcus. Finally, the medial limit was designated as the inter-hemispheric sulcus.

Removed in the revised version: Boundaries of the dorso-lateral prefrontal cortex

Earlier version: a)for great apes and humans. From left to right: sagittal view, external view, coronal view. Abbreviations: cs: cingulate sulcus; cc: corpsus callosum; frsi: frontal inferior sulcus; frss: frontal superior sulcus; cent.s: central sulcus; prec.s: precentral sulcus; A: anterior: P: posterior; D: dorsal; V: ventral; L: lateral; M: medial.

Revised version: Boundaries of the dorso-lateral prefrontal cortex a) for great apes and humans. From left to right: sagittal view, external view, coronal view. Abbreviations: cs: cingulate sulcus; cc: corpsus callosum; frsi: frontal inferior sulcus; frss: frontal superior sulcus; cent.s: central sulcus; prec.s: precentral sulcus; A: anterior: P: posterior; D: dorsal; V: ventral; L: lateral; M: medial. b) for monkeys. From left to right: sagittal view, external view, coronal view. Abbreviations: cs: cingulate sulcus; cc: corpsus callosum; ps: principal sulcus; as: arcuate sulcus; A: anterior: P: posterior; D: dorsal; V: ventral; L: lateral; M: medial.

Removed in the revised version: b) for monkeys. From left to right: sagittal view, external view, coronal view. Abbreviations: cs: cingulate sulcus; cc: corpsus callosum; ps: principal sulcus; as: arcuate sulcus; A: anterior: P: posterior; D: dorsal; V: ventral; L: lateral; M: medial.

Socio-ecological and phylogenetic data

Eleven socio-ecological variables were selected for the analyses, gathered in different categories: body condition (body mass), diet (dietary quality index and tool use), movements and ranging behavior (daily traveled distance), social parameters (group size, population density, social system), and variables related to reproduction and life-history traits (mating system, mate guarding, seasonal breeding, and weaning age). Each variable was assessed based on the literature on wild populations, whenever possible, which was the case in a vast majority of cases (and otherwise specified, see below). We verified that these variables showed minimal correlation (see supplementary material section 2).

Earlier version: The dietary quality index (DQI) was used to characterize the richness of the dietary spectrum and was calculated from the formula DQI = 1s + 2r + 3,5a, where s is the percentage of plant structural parts in the diet, r the percentage of plant reproductive parts and a the percentage of animal preys (Sailer, Gaulin et al. 1985). Thereby, a low index (around 100) characterizes folivorous diets, while a high index (around 200) characterizes more diversified diets (including animals and fruits). Tool use represented the occurrence and complexity of using objects in feeding contexts. We took the definition of St Amant and Horton (St Amant and Horton 2008) which states that ‘tool use is the exertion of control over a freely manipulable external object (the tool) with the goal of (1) altering the physical properties of another object, substance, surface or medium (the target, which may be the tool user or another organism) via a dynamic mechanical interaction, or (2) mediating the flow of information between the tool user and the environment or other organisms in the environment.’ It includes notions of control over an object and goal. The presence or absence of tool use in feeding contexts for species in the wild was assessed from literature. Daily traveled distance was expressed in kilometers. Regarding social parameters, group size was defined as the mean social group size of a species. We also used population density (number of individuals per km2) in order to bring spatial precisions over the group size variable. Group size and population density data were collected and compiled from several primary and secondary sources (see supplementary material section 1). For group size and population density of Homo sapiens, we took an average between industrialized societies and hunter-gatherers societies. Social system was defined using the four-way categorization scheme typically used in primate studies and included solitary, pair-living, polygyny and polygynandry (see also (DeCasien, Williams et al. 2017). Mating system categories included spatial polygyny (among solitary species, agonistically powerful males defend mating access to several females), monogamy (one male is socially bonded to one breeding female), polyandry (one female is simultaneously bonded to multiple males), harem polygyny (one male is simultaneously bonded to multiple breeding females) and polygynandry (multiple males and multiple females breed within the same group, but no lasting bonds are formed) (Clutton-Brock 1989). Mate guarding is defined as a female monopolization over an extended period of time to secure paternity (Manson 1997) and is characterized by male attempts to associate and copulate with a female during the presumptive fertile period (Manson 1997, Dixson and Oxford University Press. 2012). This categorical variable was divided into two categories: species using mate guarding and those that do not usually show this behavior. Finally, in order to account for differences in lifespan between study species, the weaning age was calculated as a percentage of maximum lifespan. For some missing values from wild studies, data were taken from studies in captivity and compared to data related to close species (in the wild and in captivity). This method was applied for the weaning age and the dietary quality index of Papio papio, with other species of baboons. Moreover, the absence of tool use for Lophocebus albigena was inferred from the absence of published papers or other forms of communication on this subject.

Revised version: The dietary quality index (DQI) was used to characterize the richness of the dietary spectrum and was calculated from the formula DQI = 1s + 2r + 3,5a, where s is the percentage of plant structural parts in the diet, r the percentage of plant reproductive parts and a the percentage of animal preys (Sailer, Gaulin, Boster, & Kurland, 1985). Thereby, a low index (around 100) characterizes folivorous diets, while a high index (around 200) characterizes more diversified diets (including animals and fruits). Tool use represented the occurrence and complexity of using objects in feeding contexts. We took the definition of St Amant and Horton (St Amant & Horton, 2008) which states that “tool use is the exertion of control over a freely manipulable external object (the tool) with the goal of (1) altering the physical properties of another object, substance, surface or medium (the target, which may be the tool user or another organism) via a dynamic mechanical interaction, or (2) mediating the flow of information between the tool user and the environment or other organisms in the environment.” It includes notions of control over an object and goal. The presence or absence of tool use in feeding contexts for species in the wild was assessed from literature. Daily traveled distance was expressed in kilometers. Regarding social parameters, group size was defined as the mean social group size of a species. We also used population density (number of individuals per km2) in order to bring spatial precisions over the group size variable. Group size and population density data were collected and compiled from several primary and secondary sources (see supplementary material section 1). For group size and population density of Homo sapiens, we took an average between industrialized societies and hunter-gatherers societies. Social system was defined using the four-way categorization scheme typically used in primate studies and included solitary, pair-living, polygyny and polygynandry (see also (Alex R. DeCasien, Williams, & Higham, 2017). Mating system categories included spatial polygyny (among solitary species, agonistically powerful males defend mating access to several females), monogamy (one male is socially bonded to one breeding female), polyandry (one female is simultaneously bonded to multiple males), harem polygyny (one male is simultaneously bonded to multiple breeding females) and polygy-nandry (multiple males and multiple females breed within the same group, but no lasting bonds are formed) (Clutton-Brock, 1989). Mate guarding is defined as a female monopolization over an extended period of time to secure paternity (Manson, 1997) and is characterized by male attempts to associate and copulate with a female during the presumptive fertile period (Dixson & Oxford University Press., 2012; Manson, 1997). This categorical variable was divided into two categories: species using mate guarding and those that do not usually show this behavior. Finally, in order to account for differences in lifespan between study species, the weaning age was calculated as a percentage of maximum lifespan. For some missing values from wild studies, data were taken from studies in captivity and compared to data related to close species (in the wild and in captivity). This method was applied for the weaning age and the dietary quality index of Papio papio, with other species of baboons. Moreover, the absence of tool use for Lophocebus albigena was inferred from the absence of published papers or other forms of communication on this subject.

Earlier version: The phylogenetic tree was obtained from the 10ktrees website (https://10ktrees.nunn-lab.org/Pri-mates/downloadTrees.php, version 3). This version (Arnold, Matthews et al. 2010) provides a Bayesian inference of the primate phylogeny based on collected data for eleven mitochondrial and six autosomal genes from GenBank across 301 primate species.

Revised version: The phylogenetic tree was obtained from the 10ktrees website (https://10ktrees.nunn-lab.org/Primates/downloadTrees.php, version 3). This version (Arnold, Matthews, & Nunn, 2010) provides a Bayesian inference of the primate phylogeny based on collected data for eleven mitochondrial and six autosomal genes from GenBank across 301 primate species.

Statistical analysis

We used a phylogenetic generalized least squares (PGLS) approach to evaluate the joint influence of socio-ecological variables on the neuro-anatomical variability across species (Grafen, 1989). This approach allowed us to take into account the phylogenetic relation across species when evaluating the influence of socio-ecological variables on their neuroanatomy.

Earlier version: To identify the combination of socio-ecological variables that best predicted the size of a given brain region, given the phylogenetic relations across species, we fitted neuro-anatomical data with several PGLS models, each reflecting a specific combination of socio-ecological variables. Brain measurements and body mass were log10-transformed before analyses. For the same model several correlation structures were used: standard Brownian, Pagel’s (Pagel 1999), and OU-based (Martins and Hansen 1997) and were compared based on the smallest AIC values (Akaike 1974). For a given correlation structure, the effects of the social and ecological variables were assessed with likelihood-ratio tests. All analyses were performed with ape (Paradis and Schliep 2019).

Revised version: To identify the combination of socio-ecological variables that best predicted the size of a given brain region, given the phylogenetic relations across species, we fitted neuro-anatomical data with several PGLS models, each reflecting a specific combination of socio-ecological variables. Brain measurements and body mass were log10-transformed before analyses. For the same model, several correlation structures were used: standard Brownian, Pagel’s (Pagel, 1999), and OU-based (Martins & Hansen, 1997) and were compared based on the smallest AIC values (Akaike, 1974). For a given correlation structure, the effects of the social and ecological variables were assessed with likelihood-ratio tests. All analyses were performed with ape (Paradis & Schliep, 2019). We also calculated standardized coefficients, to facilitate comparison. Standardized coefficients were calculated by the product of the coefficient estimated by PGLS with the ratio of the standard-deviation of the predictor on the standard-deviation of the response.

Removed in the revised version: We also calculated standardized coefficients, to facilitate comparison. Standardized coefficients were calculated by the product of the coefficient estimated by PGLS with the ratio of the standarddeviation of the predictor on the standard-deviation of the response.

Because of the relatively small number of species in our sample, we assessed the reliability of the inferred models with a “leave-one-out” procedure: we removed one species from the data and the phylogenetic tree and re-fitted the model selected previously. This was repeated for each species. The details of this procedure are provided in supplementary material, section 6.

Results

Neuroanatomical measures

The average size of the regions of interest (whole brain, FP and DLPFC) are shown on figure 3. As expected, all these measures were highly correlated (all r=0.99, Pearson correlation). Not only the volumes of DLPFC and FP, which are two neighboring regions of the prefrontal cortex, but also each of these regions and the whole brain.

Average values of the 3 regions of interest. Each line provides the cumulated volumes of the dorso-lateral prefrontal cortex (DLPFC, green), the frontal pole (FP, grey) and the rest of the brain (ROB, blue), such that the size of each bar represents the whole brain volume. Note: for Homo sapiens, the bar has been truncated, since its value was out of scale with the other species.

Influence of socio-ecological variables on whole brain volume

We used a model-comparison approach to select the best combination of socio-ecological variables accounting for the variability in whole brain size across species using PGLS. Details of the different models are provided in supplementary material (SI section 3).

The best model explaining the volume of the whole brain is the one that includes body mass, daily traveled distance, and population density. The values of the estimated coefficients are provided in table 1. These three variables have a positive influence on the volume of the whole brain. Even if all three factors had a significant effect, their influence on the whole brain volume seemed to differ: there was an order of magnitude between the estimated beta (i.e. contribution) of body mass (0.47 +/-0.15) and that of daily traveled distance (0.05 +/-0.01). There was also an order of magnitude between the estimated beta of daily traveled distance and population density (0.007 +/-0.003). Note that the standardized beta (see methods) is also greater for daily traveled distance than for population density (see supplementary material section 5).

We evaluated the robustness of the influence of these socio-ecological variables on whole brain volume using a « leave-one-out » procedure. Thus, each model was tested 16 times (one per species removed) and we evaluated the reliability of each variable after each species was removed. The details of this analysis are provided in the supplementary material (section 4). The influence of body mass was very robust, in that it remained significant for all models where one species was removed. Daily traveled distance remained significant for all models but one, whereas population density lost significance (p>0.05) for four models where one species was removed, even if in all cases the estimated coefficient remained positive. Thus, the influence of body mass and daily traveled distance are very robust for the combination of species used for the analysis. It is less the case for population density, even if based on model comparison this variable plays a significant role in increasing the volume of the whole brain.

In summary, this analysis indicates that the volume of the whole brain across primates is positively modulated by body mass, as well as social (population density) and ecological (daily traveled distance) variables. The weight of the ecological variable, however, seems to be slightly stronger and more reliable than the weight of the social variable.

Influence of socio-ecological variables on Frontal Pole volume

Earlier version: As for the whole brain, the best model accounting for the volume of the frontal pole is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (section 3). The values of estimated coefficients, all positive, are provided in table 2.

Revised version: As for the whole brain, the best model accounting for the volume of the frontal pole is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (section 3). The values of estimated coefficients, all positive, are provided in table 2. Note that, as expected given the strong correlation between FP and whole brain volumes, including the whole brain volume as a covariate in the regression models canceled the effects of all other variables, such that only whole brain showed a significant effect.

As for the whole brain, the largest effect is body mass (0.69 +/-0.23), then daily traveled distance (0.07 +/-0.02) and population density (0.011 +/-0.004). By comparison with the whole brain, however, the coefficients for daily traveled distance and population density have the same order of magnitude such that their impact on the volume of the frontal pole seems to be equivalent. Note that the standardized beta for daily traveled distance is smaller than for population density, whereas it was the opposite for the whole brain, i.e. bigger for daily traveled distance compared to population density.

Earlier version: As indicated in the methods section, we also conducted an analysis using a less conservative evaluation of FP volumes in monkeys. As shown in supplementary material (section 4), the relative weight of population density decreased when adding cortical regions dorsal to the principal sulcus. Thus, when using a less conservative estimate of the FP volume, its relation with socio-ecological variables was closer to that of the whole brain.

Revised version: As indicated in the methods section, we also conducted an analysis using a less conservative evaluation of FP volumes in monkeys. As shown in supplementary material (section 4), the relative weight of population density decreased when adding cortical regions dorsal to the principal sulcus.

Added in the revised version: Thus, when using a less conservative estimate of the FP volume, its relation with socio-ecological variables was closer to that of the whole brain.

As for the whole brain, we evaluated the reliability of the model’s variables using a “leave-one-out” procedure at the species level. The results of this analysis were similar to what we reported for the whole brain: the influence of body mass remained significant whatever the species removed, only one species caused daily travelled distance to lose significance and three species caused population density to lose significance (p>0.05). As can be seen in the supplementary material (section 4), however, the decrease in the weight of the estimated coefficient for population density was less dramatic for these three species, compared to what we observed for the whole brain. As specified in the method section, we also conducted an analysis where the dorsal limit of the frontal pole was less conservative, to include a part of the brain which might belong to the FP but also to the DLPFC (see results in supplementary material, section 4).

Altogether, these data indicate that the volume of the frontal pole across primates is positively modulated by body mass as well as by both population density and daily traveled distance, in line with the idea that it is affected by both social and ecological components. By comparison with the whole brain, the relative influence of the social (Population density) and ecological (Daily Traveled Distance) variables are more balanced.

Influence of socio-ecological variables on Dorso-Lateral Prefrontal Cortex volume

Earlier version: As for the whole brain and the frontal pole, the best model for the DLPFC is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (section 3). The values of estimated coefficients, all positive, are provided in table 3.

Revised version: As for the whole brain and the frontal pole, the best model for the DLPFC is the one that includes body mass, daily traveled distance and population density. Details of the different models are provided in supplementary material (section 3). The values of estimated coefficients, all positive, are provided in table 3. As it was the case for the FP, including the whole brain volume as a covariate in the regression models canceled the effects of all other variables, such that only whole brain showed a significant effect.

As for the whole brain and the frontal pole, the largest effect is body mass (0.47 +/-0.18), followed by daily traveled distance (0.047 +/-0.01). Even if the best model for DLPFC includes population density, its influence was very small (0.006 +/-0.003) and failed to reach significance (p=0.067). The « leave-one-out » procedure confirmed the reliability of the model for body mass (with no species causing that variable to lose significance when removed) and that of daily traveled distance (only one species causing that variable to lose significance when removed). The influence of population density only reached significance for two of the models where one species was removed, but it failed to reach significance for all other combinations tested (n=14), as it was the case for the original one with all 16 species. Details of this analysis can be found in the supplementary material, section 4. In other words, the relative weakness of the influence of population density compared to that of other variables is relatively reliable across the combinations of species used for the tests.

In summary, the volume of the DLPFC across primates is also positively modulated by body mass as well as by both ecological (daily traveled distance) and social (population density) variables, but the influence of the later appears much weaker than that of the former.

Discussion

Earlier version: We evaluated the influence of eleven socio-ecological variables on the size of the whole brain as well as two specific brain regions involved in executive functions: the frontal pole and the dorsolateral prefrontal cortex. All together, our results showed that all three cerebral measures strongly correlated and were influenced by the same set of socio-ecological variables: body mass, daily traveled distance and population density. As expected from evidence suggesting brain-body covariation (Martin 1981), body mass had a strong and reliable influence on all brain measures. Daily traveled distance, a proxy for how challenging is foraging, also had a clear positive influence on all brain measures. Finally, population density, a proxy for how challenging and complex are social interactions, was generally less powerful at explaining neuro-anatomical variability, but it was also the variable showing the greatest difference in effect size across brain regions. For the FP, the influence of population density was similar to that of daily traveled distance, but it was much weaker for the whole brain, and it even failed to reach significance for the DLPFC. Thus, our data are generally compatible with the idea that the evolution of executive functions relying upon FP and DLPFC is driven both by ecological and social constraints. Critically, however, the relative influence of these constraints seems to vary across regions, in line with our hypothesis based on the cognitive functions with which they are associated in laboratory conditions. Thus, neuro-cognitive entities as defined in laboratory conditions can readily be articulated with socio-ecological processes in the wild, in order to better understand primate cognition and behavior in natural conditions.

Revised version: We evaluated the influence of eleven socio-ecological variables on the size of the whole brain as well as two specific brain regions involved in executive functions: the frontal pole and the dorsolateral prefrontal cortex. Altogether, our results showed that all three cerebral measures strongly correlated and were influenced by the same set of socio-ecological variables: body mass, daily traveled distance and population density. As expected from evidence suggesting brain-body covariation (Martin, 1981), body mass had a strong and reliable influence on all brain measures. Daily traveled distance, a proxy for how challenging is foraging, also had a clear positive influence on all brain measures. Finally, population density, a proxy for how challenging and complex are social interactions, was generally less powerful at explaining neuro-anatomical variability, but it was also the variable showing the greatest difference in effect size across brain regions. For the FP, the influence of population density was similar to that of daily traveled distance, but it was much weaker for the whole brain, and it even failed to reach significance for the DLPFC. Thus, our data are generally compatible with the idea that the evolution of executive functions relying upon FP and DLPFC is driven both by ecological and social constraints. Critically, the relative influence of these constraints seems to vary across regions, in line with our hypothesis based on the cognitive functions with which they are associated in laboratory conditions. Thus, neuro-cognitive entities as defined in laboratory conditions can readily be articulated with socio-ecological processes in the wild, in order to better understand primate cognition and behavior in natural conditions.

Removed in the revised version: First, we need to consider two methodological questions:

Removed in the revised version: How reliable is the difference we reported between brain regions? Even if the same combination of variables accounted for the relative size of all three brain measures, there was a difference in the weight of population density, proxy for social challenges, across the three regions. Indeed, for the FP, the effect of population density was as strong as that of daily traveled distance: both were significant, and their effect size (beta weight) was of the same order of magnitude. Also, they showed a similarly small sensitivity to the « leave-one-out » procedure. Finally, standardized beta weights for population density were greater than for daily traveled distance. By contrast, for the DLPFC, the effect of population density failed to reach significance (p=0.07), its beta weight was one order of magnitude smaller than that of daily traveled distance and standardized beta weights for population density were smaller than for daily traveled distance. Finally, this relatively weak influence of population density on DLPFC volume was confirmed by the ‘leave-one-out’ procedure, since the effect of population density on the size of the DLPFC was reliably marginal across all combination of species used to fit the model. Interestingly, the whole brain volume (which is strongly correlated with both FP and DLPFC) seems to show an intermediate tendency: even if the effect of population density appears smaller than that of daily traveled distance (beta weight is one order of magnitude smaller), it is clearly significant, and the « leave-one-out » procedure indicated that it was only slightly less reliable than daily traveled distance, with a significant decrease in model fit quality for four vs. one combination of species. Altogether, this indicates that even if the influence of population density cannot be ruled out for any of the three brain measures, it appears quantitatively smaller for the DLPFC than for the FP. Note that the relative weight of population density also decreased (relative to that of daily traveled distance) when we used a less conservative estimate of FP volumes, which included a dorsal part potentially overlapping with the DLPFC. This is also compatible with the fact that the influence of population density seems to be of intermediate magnitude in the whole brain, which includes both FP and DLPFC.

Removed in the revised version: To what extent do these conclusions depend upon the specific sample of species used here? The results of the « leave-one-out » analysis showed that our sample, even if relatively limited, was sufficient to evaluate the relative weight of socio-ecological variables on specific brain regions in primates. Interestingly, the relation between brain region volumes and socio-ecological variables was not completely independent from the sample of species, as previously shown in a study conducted on a larger sample (Powell, Isler et al. 2017). In line with the main analysis (PGLS with all species included), the influence of body mass and daily traveled distance appears very reliable (no more than one species caused the model to fail when removed). The influence of population density appeared less reliable, since it was sensitive to the removal of up to three-four species (as a function of brain region, see previous paragraph). But the influence of individual species is nearly impossible to interpret, because no apparent obvious pattern emerges. Indeed, the species without which the model failed was not always the same (e.g. humans and baboons have very distinct values for population density, 56 vs. 7.5). Note, however, that getting clear intuitions from the data remains difficult because the PGLS includes the covariance induced by phylogeny to the linear relation across variables. We see no reason to exclude humans (as done in previous studies e.g. (DeCasien and Higham 2019), because even if the brain of Homo sapiens is much bigger than others, a general biological law describing the relation between socio-ecology and the brain in primates should apply to humans (Gabi, Neves et al. 2016, Herculano-Houzel 2017). Moreover, even if the human brain is much bigger than that of other primates, the global organization of the frontal cortex appears qualitatively similar in humans and other primates (Sallet, Mars et al. 2013, Gabi, Neves et al. 2016, Herculano-Houzel, Kaas et al. 2016, Barrett, Dawson et al. 2020, Roumazeilles, Eichert et al. 2020). One potential issue with humans is the difficulty to evaluate socio-ecological variables, given that modern human populations show a tremendous geographical variability in terms of socio-ecological variables. But the same problematic applies to other modern primate species (including those confronted to intense anthropization of their habitat) (Beauchamp and Cabana 1990). Thus, the actual potential limitation (which concerns all species) is the reliability with which socio-ecological variables were estimated for each species, given the amount of intraspecific variability. Note, there is also a significant amount of intra-specific variability at the level of the brain, both in humans and non-human primates, but this intra-specific variability was shown to be negligible compared to inter-specific variability (Maguire, Gadian et al. 2000, DeCasien and Higham 2019, Louail, Gilissen et al. 2019, Testard, Brent et al. 2022). Thus, we acknowledge that intra-specific variability could introduce noise in the inter-specific relation between brain measures and socio-ecological variables, but the fact that clear relations could be established with our sample indicates that this source of intra-specific variability was limited enough relative to the interspecies relation between neuroanatomical measures and socio-ecological variables. In other words, the error with which these variables were estimated at the level of individual species was small enough to allow us to study the relation of interest across species, and from that perspective, there is no reason to exclude any species from that analysis and our sample appears reliable enough to characterize the relation between neuroanatomical features and socio-ecological variables across primates.

Removed in the revised version: Our results provide an original insight on the existing debate on the source of inter-species variability in cognitive abilities, as measured using brain volumes, and on the main socio-ecological variables that affect the evolution of specific brain regions. Given that all three brain measures were highly correlated, it is not surprising that their variability across species is predicted by the same model, i.e. the same combination of socio-ecological variables. This is in line with the intuition that across primates the major difference in brain volumes is often accounted for by a difference in size of the prefrontal cortex (Passingham, Passingham et al. 2012); but see (DeCasien, Barton et al. 2022). Importantly, more precise quantitative studies showed that the fraction of neurons in the prefrontal cortex (PFC), relative to the whole brain (8%) was well conserved across primates, such that the key difference across species is not the proportion of prefrontal neurons, but its absolute number (Herculano-Houzel, Collins et al. 2007). In other words, the strong development of executive functions in species with larger prefrontal cortices is related to an absolute increase in number of neurons in the PFC, rather than in an increase in the ratio between the number of neurons in the PFC vs the rest of the brain. In that frame, it seems more appropriate to use absolute measures rather than relative measures, to evaluate the weight of each brain region and its corresponding cognitive function on behavior.

Removed in the revised version: In addition, the exact boundary between FP and DLPFC as functional entities remains difficult to assess, even if the functions attributed to these two regions are clearly distinct (Gallagher and Frith 2003, Passingham and Sakai 2004, Boorman, Behrens et al. 2009, Genovesio, Wise et al. 2013, Koechlin 2016, Preuss and Wise 2021). Our aim here was clearly not to provide a clear identification of anatomical boundaries across brain regions in individual species, as others have done using much finer neuro-anatomical methods. Such a fine neuro-anatomical characterization appears impossible to carry on for a sample size of species compatible with PGLS. More importantly, we make no claim about specific mapping between the regions measured here and either the specific cognitive operations of interest (metacognition vs working memory) or the socio-ecological variables (e.g. social or foraging complexity). First, we are aware of the complexity of identifying neural processes underlying such high level cognitive functions, and even in well controlled laboratory conditions the exact boundaries of cortical regions remains challenging (Petrides 2005, Sallet, Mars et al. 2013, Mansouri, Koechlin et al. 2017). Rather, as specified in the methods, we tried to maximize the reliability of the landmarks that could be identified in all species to compare the relative size of areas that are essentially used as proxies for their known function, as defined in laboratory conditions in macaques and humans. The underlying assumption that we used was that the functional anatomy of the prefrontal cortex was conserved enough in primates such that traits that were common between humans and macaques were also shared with other primates (Sallet, Mars et al. 2013, Bludau, Eickhoff et al. 2014, Amiez, Sallet et al. 2019). Note that studies conducted in marmosets confirm that the functional anatomy of the primate prefrontal cortex is relatively well conserved (Dias, Robbins et al. 1996, Roberts, Tomic et al. 2007). Of course, the neural mechanisms underlying metacognition and working memory are extremely complicated and they involve myriads of other brain regions (e.g. temporal and parietal cortices, as well as subcortical systems), but the contribution of the FP to metacognition and of the DLPFC to working memory is so strong and reliable in humans and laboratory primates that it seems reasonable to assume that, given the relative level of conservation of prefrontal cortex anatomical organization across primates, FP and DLPFC should also play a key role in metacognition and working memory in other primate species, at least those studied here. As for our previous study using a similar approach, this method was reliable enough to capture meaningful and specific effects of interests (Louail, Gilissen et al. 2019). Thus, even if we make no strong claim regarding the specific boundaries of FP and DLPFC, the anatomo-functional difference between these regions, as characterized in a few species used in laboratory experiments is reliable enough to be captured by the variability in socio-ecological variables in a larger set of primates.

Earlier version: In summary, we are aware of the limitations of our study in terms of anatomical specificity, but these measures appear robust enough to draw original conclusions regarding the relation between functions attributed to FP and DLPFC (metacognition and working memory) in laboratory conditions and socio-ecological variables across 16 primate species. As we and others have done before (within and across species), we assume that the size of a given brain region provides a good proxy of the strength of its influence (through its known function) on behavior (Barton and Harvey 2000, Maguire, Gadian et al. 2000, DeCasien and Higham 2019, Louail, Gilissen et al. 2019). Again, we assume that the size of a brain region is a good proxy for the number of neurons allocated to its function, given the known positive relation between number of cortical neurons and cognitive skills in primates (Herculano-Houzel 2017, Herculano-Houzel 2018). Here, we restricted our analysis to a set of specific functional regions, namely the FP and the DLPFC, which implies that a given region and the neurons that it contains are characterized by a relatively homogenous connectivity pattern and physiological features, as demonstrated by converging evidences from anatomy, physiology, imaging and behavioral tests in laboratory animals (see methods for further details). In that frame, the difference in size of these brain regions across species could readily be used as proxies for the relative importance of their corresponding functions (namely, metacognition for the FP and working memory for the DLPFC). Again, we clearly make no strong claim regarding the exact boundaries of the regions of interest, and their relations with existing laboratory studies in humans or macaques. Rather, we wanted to assess the relation between neuro-cognitive entities (e.g. FP/metacognition vs. DLPFC/working memory) and socio-ecological variables, based on the assumption that the more a given cognitive function would be required to face a given socio-ecolo-gical challenge in a given species, the bigger the corresponding brain region would be. From that perspective, these data provide a new insight into the function of the FP and the DLPFC in primates.

Revised version: As we and others have done before (within and across species), we assume that the size of a given brain region provides a good proxy of the strength of its influence (through its known cognitive function) on behavior (Barton & Harvey, 2000; Alex R. DeCasien & Higham, 2019b; Louail et al., 2019; Maguire et al., 2000). Indeed, the size of a brain region provides a reliable estimate of the number of neurons allocated to its function, given the known positive relation between number of cortical neurons and cognitive skills in primates (Herculano-Houzel, 2017, 2018). Here, we restricted our analysis to a set of specific brain regions, namely the FP and the DLPFC, assuming that the difference in size of these brain regions across species could readily be used as proxies for the relative importance of their corresponding functions (namely, metacognition for the FP and working memory for the DLPFC). Using that tool, we explored the relation between neuro-cognitive entities (e.g. FP/metacognition vs. DLPFC/working memory) and socio-ecological variables, based on the assumption that the more a given cognitive function would be required to face a given socio-ecological challenge, the bigger the corresponding brain region would be.

Added in the revised version: Our aim here was clearly not to provide an accurate identification of anatomical boundaries across brain regions in individual species, as others have done using finer neuro-anatomical methods in humans and macaques (e.g. Petrides et al, 2012). Moreover, we cannot affirm that our conclusions in gyrencephalic species could be extented to lissencephalic species. Indeed, we acknowledge that since the definition of the regions of interest we used is based on gyri and sulci, it was therefore impossible to include lissencephalic primates in our sample, even if previous laboratory studies in marmosets could identify DLPFC and FP using both cytoarchitectonic and functional criteria (Dias, Robbins, & Roberts, 1996; Dureux, Zanini, Selvanayagam, Menon, & Everling, 2023; Roberts et al., 2007; Wong, Selvanayagam, Johnston, & Everling, 2023). Nonetheless, since our comparative study and the associated phylogenetic analysis required at least 10-15 species, it could not have been conducted with only laboratory species and these invasive approaches, and we therefore focused on gyrencephalic species.

Added in the revised version: The mapping between the regions measured here and either the specific cognitive operations of interest (metacognition vs working memory) or the socio-ecological variables (e.g. social or foraging complexity) is probably not exclusive. First, we are aware of the complexity of identifying neural processes underlying such high level cognitive functions, and even in well-controlled laboratory conditions, the delimitation of exact boundaries of cortical regions remains challenging (Mansouri et al., 2017; M. Petrides, 2005; J. Sallet et al., 2013). Rather, we tried to maximize the reliability of the landmarks that could be identified in all species to compare the relative size of areas that are essentially used as proxies for their known function, as defined in laboratory conditions in macaques and humans. The underlying assumption that we used was that the functional anatomy of the prefrontal cortex was enough conserved in primates, such that traits that were common between humans and macaques were also shared with other primates (Amiez et al., 2019; Bludau et al., 2014; J. Sallet et al., 2013). Of course, the neural mechanisms underlying metacognition and working memory are extremely complicated and they involve myriads of other brain regions (e.g. temporal and parietal cortices, as well as subcortical systems). Moreover, the exact boundary between FP and DLPFC as functional entities remains difficult to assess, even in laboratory conditions (Boorman, Behrens, Woolrich, & Rushworth, 2009; Gallagher & Frith, 2003; Genovesio, Wise, & Passingham, 2013; Koechlin, 2016; Passingham & Sakai, 2004; Preuss & Wise, 2021). But the contribution of the FP to metacognition and of the DLPFC to working memory is so strong and reliable in humans and laboratory primates that it seems reasonable to assume that, given the relative level of conservation of prefrontal cortex anatomical organization across primates, FP and DLPFC should also play a key role in metacognition and working memory in other primate species. As for our previous study using a similar approach, this method was reliable enough to capture meaningful and specific effects of interest (Louail et al., 2019). Thus, even if we make no strong claim regarding the specific boundaries of FP and DLPFC, the anatomo-functional difference between these regions, as characterized in a few species used in laboratory experiments, is reliable enough to be captured by the variability in socio-ecological variables in a larger set of primates.

Added in the revised version: In previous studies, the specificity of the relation between a given brain region and socioecological variables has often been evaluated using relative measures, such that the variable represents the relative increase in volume of the region relative to the rest of the brain (e.g. Krebs et al, 1989; Barton et al, 1995). This approach enables to detect changes in relation to allometric relations, which are often interpreted in terms of significant evolutionary events (major reorganization of the brain). By contrast, changes in regional volume that remain within allometric proportions (in the brain) are interpreted as a non-specific effect, and therefore negligible from an evolutionary perspective (Barton & Harvey, 2000; Barton et al., 1995; A. R. DeCasien et al., 2022; Krebs et al., 1989; Smaers et al., 2017). But even if we share the intuition of specificity captured by relative measures, we decided to use absolute measures because relative measures have several theoretical implications that we find problematic. First, relative measures imply that the effect of interest should be negligible in the rest of the brain, compared to the area of interest. But we made no predictions about the rest of the brain given that most cognitive operations involve networks rather than unique regions. For example, since working memory also relies upon parietal cortices, we expect parietal cortices to also show a positive relation with ecological constraints. Thus, the theoretical effect size of a relative measure, given the functional specificity hypothesis, would require identifying the relative volume of all regions involved in working memory (for DLPFC) and metacognition (for FP), which is clearly beyond the scope of this study. As it will be discussed in the next paragraph, we believe that the comparison between DLPFC, FP and whole brain measures provides a good index of specificity of the effects, and they can be interpreted more directly. Second, the interpretation of relative measures implies that, all other things being equal, a pool of neurons involved in a given function is less efficient when surrounded by a large number of neurons, compared to the situation where the same pool is surrounded by fewer neurons. For example, the hippocampus occupies a much larger fraction of the brain in rodents compared to primates, such that based on a relative size measurement, one would conclude that rodents have better spatial and episodic memory skills compared to primates. And even if the comparison remains difficult, there is clearly no evidence for that in behavioral data (Crystal, 2009; de Cothi et al., 2022). In summary, even if we acknowledge that relative measures provide an index of the specificity of the relation between socio-ecological variables and regional brain volumes, we find relative measures much more difficult to interpret and we chose to use absolute measures of cerebral volumes to stay away from these considerations. Indeed, the number of neurons in a given brain region should increase if its function can promote survival, irrespectively of what happens in other regions.

Added in the revised version: We did evaluate the specificity of the effects by comparing the influence of socio-ecological variables not only between the two specific prefrontal regions but also with the whole brain volume. At first sight, the very strong correlation between the volume of the whole brain and that of both FP and DLPFC implies that the specificity of the influence of socio-ecological factors on these prefrontal regions should be limited, and indeed, the variability of all 3 measures (FP, DLPFC and whole brain) is predicted by the same model, i.e. the same combination of socio-ecological variables. This is in line with the idea that executive functions are strongly inter-related and it can prove difficult to demonstrate their independence, even within a single species in controlled conditions (Miyake et al., 2000; Volter et al., 2022). But in spite of this strong correlation and in line with our prediction based on laboratory studies, we could still find significant differences between FP and DLPFC in terms of relation with population density, an index of social complexity.

Added in the revised version: How reliable is the difference we reported between brain regions? Even if the same combination of variables accounted for the relative size of all three brain measures, there was a difference in the weight of population density, proxy for social challenges, across the three regions. Indeed, for the FP, the effect of population density was as strong as that of daily traveled distance: both were significant, and their effect size (beta weight) was of the same order of magnitude. Also, they showed a similarly small sensitivity to the « leave-one-out » procedure. Finally, standardized beta weights for population density were greater than for daily traveled distance. By contrast, for the DLPFC, the effect of population density failed to reach significance (p=0.07), its beta weight was one order of magnitude smaller than that of daily traveled distance and standardized beta weights for population density were smaller than for daily traveled distance. Finally, this relatively weak influence of population density on DLPFC volume was confirmed by the “leave-one-out” procedure, since the effect of population density on the size of the DLPFC was reliably marginal across all combination of species used to fit the model. Interestingly, the whole brain volume (which is strongly correlated with both FP and DLPFC) seems to show an intermediate tendency: even if the effect of population density appears smaller than that of daily traveled distance (beta weight is one order of magnitude smaller), it is clearly significant, and the « leave-one-out » procedure indicated that it was only slightly less reliable than daily traveled distance, with a significant decrease in model fit quality for four vs. one combination of species. Altogether, this indicates that even if the influence of population density cannot be ruled out for any of the three brain measures, it appears quantitatively smaller for the DLPFC than for the FP. Note that the relative weight of population density also decreased (relative to that of daily traveled distance) when we used a less conservative estimate of FP volumes, which included a dorsal part potentially overlapping with the DLPFC. This is also compatible with the fact that the influence of population density seems to be of intermediate magnitude in the whole brain, which includes both FP and DLPFC. Again, we acknowledge that this method has limitations and that further studies, including more species, would be necessary to evaluate the nature and the dynamics of underlying evolutionary processes.

Added in the revised version: To what extent do these conclusions depend upon the specific sample of species used here? The results of the « leave-one-out » analysis showed that our sample, even if relatively limited, was sufficient to evaluate the relative weight of socio-ecological variables on specific brain regions in primates. Interestingly, the relation between brain region volumes and socio-ecological variables was not completely independent from the sample of species, as previously shown in a study conducted on a larger sample (Powell et al., 2017). In line with the main analysis (PGLS with all species included), the influence of body mass and daily traveled distance appears very reliable (no more than one species caused the model to fail when removed). The influence of population density appeared less reliable, since it was sensitive to the removal of up to three-four species (as a function of brain region, see previous paragraph). But the influence of individual species is nearly impossible to interpret, because no apparent obvious pattern emerges. Indeed, the species without which the model failed was not always the same (e.g. humans and baboons have very distinct values for population density, 56 vs. 7.5). Note, however, that getting clear intuitions from the data remains difficult because the PGLS includes the covariance induced by phylogeny to the linear relation across variables. We see no reason to exclude humans (as done in previous studies e.g. (Alex R. DeCasien & Higham, 2019b), because even if the brain of Homo sapiens is much bigger than others, a general biological law describing the relation between socio-ecology and the brain in primates should also apply to humans (Gabi et al., 2016; Herculano-Houzel, 2017). Moreover, even if the human brain is much bigger than that of other primates, the global organization of the frontal cortex appears qualitatively similar in humans and other primates (Barrett et al., 2020; Gabi et al., 2016; Herculano-Houzel et al., 2016; Roumazeilles et al., 2020; J. Sallet et al., 2013). One potential issue with humans is the difficulty to evaluate socio-ecological variables, given that modern human populations show a tremendous geographical variability in terms of socio-ecological variables. But the same problematic applies to other modern primate species (including those confronted to intense anthropization of their habitat) (McKinney, Waters, & Rodrigues, 2023). Thus, the actual potential limitation (which concerns all species) is the reliability with which socio-ecological variables were estimated for each species, given the amount of intra-specific variability. Note, there is also a significant amount of intra-specific variability at the level of the brain, both in humans and non-human primates, but this intra-specific variability was shown to be negligible compared to inter-specific variability (Alex R. DeCasien & Higham, 2019a; Louail et al., 2019; Maguire et al., 2000; Testard et al., 2022). Thus, we acknowledge that intra-specific variability could introduce noise in the inter-specific relation between brain measures and socio-ecological variables, but the fact that clear relations could be established with our sample indicates that this source of intra-specific variability was limited enough relative to the inter-species relation between neuroanatomical measures and socio-ecological variables. In other words, the error with which these variables were estimated at the level of individual species was small enough to allow us to study the relation of interest across species, and from that perspective, there is no reason to exclude any species from that analysis and our sample appears reliable enough to characterize the relation between neuroanatomical features and socio-ecological variables across primates.

Earlier version: The stronger sensitivity of the frontal pole to the variable ‘population density’ is reminiscent of the social brain hypothesis, i.e. the idea that social challenges favored the evolution of larger brains, and especially larger neocortex size, with species living in larger groups having bigger brains to deal with the associated complex social interactions (Dunbar 1998). Critically, the size of the frontal pole and the corresponding development of metacognitive skills is not exclusively related to population density (proxy for social interactions) but also to daily traveled distance (proxy for foraging complexity). Thus, these data are also compatible with the ecological brain hypothesis (Milton 1981). This dual relation between frontal pole size, social and ecological constraints might be accounted for by the fact that in the wild, social and ecological constraints remain strongly related, even if we did not find strong correlations between those factors in our data set. In reality, it is difficult to treat ecological and social factors as if they were disconnected (Henke-von der Malsburg, Kappeler et al. 2020), and cognitive abilities associated with foraging are likely to play a role in social foraging tactics too (Street, Navarrete et al. 2017). Indeed, greater population density implies a greater inter-individual competition for food, and thus potentially increases in daily traveled distances, and/or the development of sophisticated foraging skills to deal with the increase in scramble competition. But from a cognitive point of view, this also implies that the benefits associated with increased FP volume, and presumably an increase in metacognitive skills, could be related to both social and ecological functions. Indeed, as pointed out earlier, the development of executive functions in general, and metacognitive skills in particular, could have been a critical leverage to allow the development of both complex social interactions (through the use of Theory of Mind) and complex foraging (through flexible, context-dependent planning) (Garcia, Bouret et al. 2021, Shultz and Dunbar 2022). Thus, our work indicates that the notions of ‘social brain’ and ‘ecological brain’ should not be mutually exclusive, even when considering specific brain regions. Rather, we confirm intuitions based on laboratory studies that the frontal pole, through its role in metacognition, belongs to both sets of brain regions defined as the social brain and the ecological brain, respectively. Further work would be needed to capture the metacognitive processes underlying social and foraging processes in the wild, as well as their interactions.

Revised version: The stronger sensitivity of the frontal pole to the variable ‘population density’ is reminiscent of the social brain hypothesis, i.e. the idea that social challenges favored the evolution of larger brains, and especially larger neocortex size, with species living in larger groups having bigger brains to deal with the associated complex social interactions (Dunbar, 1998). Critically, the size of the frontal pole and the corresponding development of metacognitive skills is not exclusively related to population density (proxy for social interactions) but also to daily traveled distance (proxy for foraging complexity). Thus, these data are also compatible with the ecological brain hypothesis (Milton, 1981). This dual relation between frontal pole size, social and ecological constraints might be accounted for by the fact that in the wild, social and ecological constraints remain strongly related, even if we did not find strong correlations between those factors in our data set. In reality, it is difficult to treat ecological and social factors as if they were disconnected (Henke-von der Malsburg, Kappeler, & Fichtel, 2020), and cognitive abilities associated with foraging are likely to play a role in social foraging tactics too (Street, Navarrete, Reader, & Laland, 2017). Indeed, greater population density implies a greater inter-individual competition for food, and thus potentially increases in daily traveled distances, and/or the development of sophisticated foraging skills to deal with the increase in scramble competition. But from a cognitive point of view, this also implies that the benefits associated with increased FP volume, and presumably an increase in metacognitive skills, could be related to both social and ecological functions. Indeed, as pointed out earlier, the development of executive functions in general, and metacognitive skills in particular, could have been a critical leverage to allow the development of both complex social interactions (through the use of Theory of Mind) and complex foraging (through flexible, context-dependent planning) (Garcia et al., 2021; Shultz & Dunbar, 2022). Thus, our work indicates that the notions of ‘social brain’ and ‘ecological brain’ should not be mutually exclusive, even when considering specific brain regions. Rather, we confirm intuitions based on laboratory studies that the frontal pole, through its role in metacognition, belongs to both sets of brain regions defined as the social brain and the ecological brain, respectively. Further work would be needed to capture the metacognitive processes underlying social and foraging processes in the wild, as well as their interactions.

Earlier version: The size of the DLPFC, which we used as a proxy for working memory and planning skills, showed a significant relation with daily traveled distance (proxy for foraging complexity), and to a lesser extent with population density (proxy for social complexity). This suggests that the cognitive functions at play in the DLPFC, i.e. working memory and planning, are critical for foraging in primates, especially when they need to travel long distances, and presumably have to deal with more complex navigation strategies. This is clearly in line with the global idea of the ecological brain hypothesis, but here we provide a critical insight into the specific neuro-cognitive operations associated with foraging strategies. The weaker influence of population density, a marker of social complexity, might be surprising at first sight, because a priori working memory and planning could also be strongly involved in complex social interactions (Garcia, Bouret et al. 2021). But it is in line with laboratory data showing that social interactions seem to rely much more upon rostro-medial prefrontal regions compared to DLPFC (Frith 2007, Sallet, Mars et al. 2011, Fleming and Dolan 2012, Testard, Brent et al. 2022). In other words, this work confirms the specificity of the ‘social brain’’ as defined in laboratory conditions as a set of structures specifically involved in social interactions (Frith 2007, Rushworth, Mars et al. 2013). Critically, again, social and ecological functions are tightly intermingled in primates’ natural environment and more specific studies would be needed to clarify how DLPFC-related functions such as working memory are involved in natural conditions, when animals need to face both ecological and social challenges.

Revised version: The size of the DLPFC, which we used as a proxy for working memory and planning skills, showed a significant relation with daily traveled distance (proxy for foraging complexity), and to a lesser extent with population density (proxy for social complexity). This suggests that the cognitive functions at play in the DLPFC, i.e. working memory and planning, are critical for foraging in primates, especially when they need to travel long distances, and presumably have to deal with more complex navigation strategies. This is clearly in line with the global idea of the ecological brain hypothesis, but here we provide a critical insight into the specific neuro-cognitive operations associated with foraging strategies. The weaker influence of population density, a marker of social complexity, might be surprising at first sight, because a priori working memory and planning could also be strongly involved in complex social interactions (Garcia et al., 2021). But it is in line with laboratory data showing that social interactions seem to rely much more upon rostro-medial prefrontal regions compared to DLPFC (Fleming & Dolan, 2012; Frith, 2007; Jérôme Sallet et al., 2011; Testard et al., 2022). In other words, this work confirms the specificity of the ‘social brain’’ as defined in laboratory conditions as a set of structures specifically involved in social interactions (Frith, 2007; Rushworth et al., 2013). Critically, again, social and ecological functions are tightly intermingled in primates’ natural environment and more specific studies would be needed to clarify how DLPFC-related functions such as working memory are involved in natural conditions, when animals need to face both ecological and social challenges.

Earlier version: This new set of results can be integrated in the general framework of the primate mosaic brain evolution, i.e. the different distinct structures varying in size both within and between species, and reflecting selection for cognitive skills. In line with recently published papers (Smaers, Rothman et al. 2021, van Schaik, Triki et al. 2021, DeCasien, Barton et al. 2022), we argue that the relation between ecology, neurobiology and cognition can be better captured with more specific brain measures, and more specific cognitive operations, than by using the traditional measures of relative brain size alone. We believe that comparisons between brain regions have the potential to identify which patterns of brain region evolution can provide insights into the evolution of cognition. Our new set of results are in agreement with a previous study we conducted on another brain region, the ventromedial prefrontal cortex (VMPFC), critically involved in value-based decision making (Louail, Gilissen et al. 2019). This study was conducted on 29 brain scans from only five species, such that we could not use a PGLS and instead included phylogenetic distance as a co-regressor, along social and ecological variables. As we did here, we also identified the combination of socioecological variables that best predicted neuro-anatomical variability across species. Interestingly, the pattern reported for the whole brain was similar to that of the current study and, the weight of the ecological variable (daily traveled distance) was one order of magnitude greater than the influence of the social variable (group size). As it was the case here, VMPFC and whole brain volumes were strongly correlated (r=0.99), but the size of the VMPFC was predicted by a distinct set of ecological variables (dietary quality and weaning age), with little modulation by group size. Thus, our results from these two combined studies suggest that specific neuro-cognitive entities, established in laboratory studies and defined by a conjunction of specific brain regions and cognitive operations, can be related to specific socio-ecological challenges that animals face in natural conditions.

Revised version: This new set of results can be integrated in the general framework of the primate mosaic brain evolution, i.e. the different distinct structures varying in size both within and between species, and reflecting selection for cognitive skills. In line with recently published papers (A. R. DeCasien et al., 2022; Smaers et al., 2021; Carel P. van Schaik et al., 2021), we argue that the relation between ecology, neurobiology and cognition can be better captured with more specific brain measures, and more specific cognitive operations, than by using the cruder measure of brain size. We believe that comparisons between brain regions have the potential to identify which patterns of brain region evolution can provide insights into the evolution of cognition. Our new set of results are in agreement with a previous study we conducted on another brain region, the ventromedial prefrontal cortex (VMPFC), critically involved in value-based decision making (Louail et al., 2019). This study was conducted on 29 brain scans from only five species, such that we could not use a PGLS and instead included phylogenetic distance as a co-regressor, along social and ecological variables. As we did here, we also identified the combination of socio-ecological variables that best predicted neuro-anatomical variability across species. Interestingly, the pattern reported for the whole brain was similar to that of the current study and, the weight of the ecological variable (daily traveled distance) was one order of magnitude greater than the influence of the social variable (group size). As it was the case here, VMPFC and whole brain volumes were strongly correlated (r=0.99), but the size of the VMPFC was predicted by a distinct set of ecological variables (dietary quality and weaning age), with little modulation by group size. Thus, our results from these two combined studies suggest that specific neuro-cognitive entities, established in laboratory studies and defined by a conjunction of specific brain regions and cognitive operations, can be related to specific socio-ecological challenges that animals face in natural conditions.

Earlier version: In conclusion, our results confirm that the size of specific brain regions can be related to socio-eco-logical variables through the cognitive operations relying on these regions in laboratory conditions. Thus, our approach aiming at articulating cognitive operations reported in laboratory settings with real socio-ecological challenges should provide a clear insight into the neuro-cognitive operations at play in the wild, as well as their evolution in primates. Conversely, integrating realistic socio-eco-logical challenges can provide a strong insight into the evolution of specific brain functions in primates. Of course, we would need to provide a clearer model regarding how facing socio-ecological challenges can rely upon specific and dynamic cognitive operations, but we believe that this is a critical first step that helps building a common theoretical framework to cross boundaries across behavioral ecology and cognitive neurosciences.

Revised version: In conclusion, our results confirm that the size of specific brain regions can be related to socio-ecological variables through the cognitive operations relying on these regions in laboratory conditions. Thus, our approach which aims at articulating cognitive operations reported in laboratory settings with real socio-ecological challenges should provide a clear insight into the neuro-cognitive operations at play in the wild, as well as their evolution in primates. Conversely, integrating realistic socio-ecological challenges can provide a strong insight into the evolution of specific brain functions in primates. Of course, we would need to provide a clearer model regarding how facing socio-ecological challenges can rely upon specific and dynamic cognitive operations, but we believe that this is a critical first step that helps building a common theoretical framework to cross boundaries across behavioral ecology and cognitive neurosciences.

Removed in the revised version: Supporting information